[1456511491] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=3.31%, iowait=59.67%, idle=35.91%, nice=0.00%, steal=1.10%
[1456511551] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=49.49% system=1.52%, iowait=0.00%, idle=48.48%, nice=0.00%, steal=0.51%
[1456511551] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=62.98% system=3.85%, iowait=0.48%, idle=32.21%, nice=0.00%, steal=0.48%
[1456511611] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=8.38% system=0.52%, iowait=0.00%, idle=91.10%, nice=0.00%, steal=0.00%
[1456531290] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 181
[1456531349] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 109
[1456531409] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 100
[1456531769] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=86
[1456537720] SERVICE ALERT: credr_stage_52.76.16.57;Ping Check;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1456537773] SERVICE ALERT: credr_stage_52.76.16.57;Ping Check;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1456555263] SERVICE ALERT: offokplese_52.74.206.204;Zombie Process;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555292] SERVICE ALERT: offokplese_52.74.206.204;CPU Usage;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555307] SERVICE ALERT: offokplese_52.74.206.204;Current Load;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555322] SERVICE ALERT: offokplese_52.74.206.204;Current Users;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555323] SERVICE ALERT: offokplese_52.74.206.204;Zombie Process;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555337] SERVICE ALERT: offokplese_52.74.206.204;Memory;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555346] SERVICE ALERT: offokplese_52.74.206.204;Total Processes;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555352] SERVICE ALERT: offokplese_52.74.206.204;CPU Usage;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555367] SERVICE ALERT: offokplese_52.74.206.204;Current Load;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555368] SERVICE ALERT: offokplese_52.74.206.204;Ping Check;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555382] SERVICE ALERT: offokplese_52.74.206.204;Current Users;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555383] SERVICE ALERT: offokplese_52.74.206.204;Root Partition;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1456555383] SERVICE ALERT: offokplese_52.74.206.204;Zombie Process;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555397] SERVICE ALERT: offokplese_52.74.206.204;Memory;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555406] SERVICE ALERT: offokplese_52.74.206.204;Total Processes;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555412] SERVICE ALERT: offokplese_52.74.206.204;CPU Usage;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555427] SERVICE ALERT: offokplese_52.74.206.204;Current Load;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555428] SERVICE ALERT: offokplese_52.74.206.204;Ping Check;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555442] SERVICE ALERT: offokplese_52.74.206.204;Current Users;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555443] SERVICE ALERT: offokplese_52.74.206.204;Root Partition;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1456555457] SERVICE ALERT: offokplese_52.74.206.204;Memory;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555466] SERVICE ALERT: offokplese_52.74.206.204;Total Processes;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555488] SERVICE ALERT: offokplese_52.74.206.204;Ping Check;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555503] SERVICE ALERT: offokplese_52.74.206.204;Root Partition;UNKNOWN;HARD;3;NRPE: Unable to read output
[1456555967] SERVICE ALERT: offokplese_52.74.206.204;Current Load;OK;HARD;3;OK - load average: 0.04, 0.03, 0.05
[1456555982] SERVICE ALERT: offokplese_52.74.206.204;Current Users;OK;HARD;3;USERS OK - 1 users currently logged in
[1456555997] SERVICE ALERT: offokplese_52.74.206.204;Memory;OK;HARD;3;[MEMORY] Total: 3952 MB - Used: 207 MB - 5% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1456556006] SERVICE ALERT: offokplese_52.74.206.204;Total Processes;OK;HARD;3;PROCS OK: 111 processes
[1456556032] SERVICE ALERT: offokplese_52.74.206.204;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1456556043] SERVICE ALERT: offokplese_52.74.206.204;Root Partition;OK;HARD;3;DISK OK - free space: / 17575 MB (92% inode=93%):
[1456556103] SERVICE ALERT: offokplese_52.74.206.204;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1456556134] SERVICE ALERT: offokplese_52.74.206.204;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456556867] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1456556912] SERVICE ALERT: motoassist_52.76.71.31;CPU Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556926] SERVICE ALERT: motoassist_52.76.71.31;CPU Usage;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556927] SERVICE ALERT: motoassist_52.76.71.31;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556934] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1456556934] SERVICE ALERT: motoassist_52.76.71.31;CPU Usage;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556942] SERVICE ALERT: motoassist_52.76.71.31;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556957] SERVICE ALERT: motoassist_52.76.71.31;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556964] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1456556987] SERVICE ALERT: motoassist_52.76.71.31;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456556988] SERVICE ALERT: motoassist_52.76.71.31;Ping Check;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1456556991] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1456557002] SERVICE ALERT: motoassist_52.76.71.31;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557002] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1456557003] SERVICE ALERT: motoassist_52.76.71.31;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557017] SERVICE ALERT: motoassist_52.76.71.31;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557032] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1456557033] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1456557034] SERVICE ALERT: motoassist_52.76.71.31;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557047] SERVICE ALERT: motoassist_52.76.71.31;Current Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557048] SERVICE ALERT: motoassist_52.76.71.31;Ping Check;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557062] SERVICE ALERT: motoassist_52.76.71.31;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557063] SERVICE ALERT: motoassist_52.76.71.31;Root Partition;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557064] SERVICE ALERT: motoassist_52.76.71.31;Zombie Process;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557064] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1456557077] SERVICE ALERT: motoassist_52.76.71.31;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557094] SERVICE ALERT: motoassist_52.76.71.31;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557107] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1456557108] SERVICE ALERT: motoassist_52.76.71.31;Ping Check;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557140] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1456557199] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1456557200] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1456557200] SERVICE ALERT: motoassist_52.76.71.31;Root Partition;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557211] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1456557215] SERVICE ALERT: motoassist_52.76.71.31;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557222] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1456557230] SERVICE ALERT: motoassist_52.76.71.31;Zombie Process;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557230] HOST ALERT: motoassist_52.76.71.31;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1456557245] HOST ALERT: motoassist_52.76.71.31;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1456557290] SERVICE ALERT: motoassist_52.76.71.31;Zombie Process;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456557480] SERVICE ALERT: credr_jenkins_52.76.31.249;Nginx http;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 1490 bytes in 0.003 second response time
[1456557540] SERVICE ALERT: credr_jenkins_52.76.31.249;Nginx http;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 1489 bytes in 0.002 second response time
[1456557600] SERVICE ALERT: credr_jenkins_52.76.31.249;Nginx http;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1490 bytes in 0.002 second response time
[1456557976] HOST ALERT: credr_wordpress_52.74.228.190;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1456558021] SERVICE ALERT: credr_wordpress_52.74.228.190;CPU Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558036] SERVICE ALERT: credr_wordpress_52.74.228.190;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558051] HOST ALERT: credr_wordpress_52.74.228.190;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1456558052] SERVICE ALERT: credr_wordpress_52.74.228.190;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558067] SERVICE ALERT: credr_wordpress_52.74.228.190;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558081] SERVICE ALERT: credr_wordpress_52.74.228.190;CPU Usage;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558082] HOST ALERT: credr_wordpress_52.74.228.190;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1456558086] SERVICE ALERT: credr_wordpress_52.74.228.190;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456558096] SERVICE ALERT: credr_wordpress_52.74.228.190;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558096] SERVICE ALERT: credr_wordpress_52.74.228.190;Current Load;OK;SOFT;3;OK - load average: 0.00, 0.02, 0.05
[1456558098] SERVICE ALERT: credr_wordpress_52.74.228.190;Ping Check;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558098] SERVICE ALERT: credr_wordpress_52.74.228.190;Memory;OK;SOFT;2;[MEMORY] Total: 3951 MB - Used: 518 MB - 13% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1456558102] SERVICE ALERT: credr_wordpress_52.74.228.190;Ping Check;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1456558112] SERVICE ALERT: credr_wordpress_52.74.228.190;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456558112] SERVICE ALERT: credr_wordpress_52.74.228.190;Current Users;OK;SOFT;3;USERS OK - 1 users currently logged in
[1456558114] HOST ALERT: credr_wordpress_52.74.228.190;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1456558130] SERVICE ALERT: motoassist_52.76.71.31;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1456558130] HOST ALERT: motoassist_52.76.71.31;UP;HARD;10;PING WARNING - Packet loss = 90%, RTA = 1.47 ms
[1456558132] SERVICE ALERT: motoassist_52.76.71.31;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1456558137] SERVICE ALERT: motoassist_52.76.71.31;Current Load;OK;HARD;3;OK - load average: 0.00, 0.01, 0.05
[1456558140] SERVICE ALERT: motoassist_52.76.71.31;Current Users;OK;HARD;3;USERS OK - 0 users currently logged in
[1456558145] HOST ALERT: credr_wordpress_52.74.228.190;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1456558155] SERVICE ALERT: motoassist_52.76.71.31;Memory;OK;HARD;3;[MEMORY] Total: 2000 MB - Used: 164 MB - 8% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1456558206] HOST ALERT: credr_wordpress_52.74.228.190;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1456558207] SERVICE ALERT: motoassist_52.76.71.31;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456558220] SERVICE ALERT: motoassist_52.76.71.31;Root Partition;OK;HARD;3;DISK OK - free space: / 12725 MB (89% inode=90%):
[1456558235] SERVICE ALERT: motoassist_52.76.71.31;Total Processes;OK;HARD;3;PROCS OK: 107 processes
[1456558411] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=54.50% system=0.50%, iowait=0.00%, idle=44.50%, nice=0.00%, steal=0.50%
[1456558471] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.48% system=0.50%, iowait=0.00%, idle=96.02%, nice=0.00%, steal=0.00%
[1456559442] SERVICE ALERT: credr_mobile_52.77.29.14;Nginx http;CRITICAL;SOFT;1;connect to address 52.77.29.14 and port 80: Connection refused
[1456559502] SERVICE ALERT: credr_mobile_52.77.29.14;Nginx http;CRITICAL;SOFT;2;connect to address 52.77.29.14 and port 80: Connection refused
[1456559562] SERVICE ALERT: credr_mobile_52.77.29.14;Nginx http;CRITICAL;HARD;3;connect to address 52.77.29.14 and port 80: Connection refused
[1456559761] SERVICE ALERT: credr_imageprocessing_52.76.65.82;CPU Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559776] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559792] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559808] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559821] SERVICE ALERT: credr_imageprocessing_52.76.65.82;CPU Usage;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559890] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Zombie Process;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559904] SERVICE ALERT: credr_imageprocessing_52.76.65.82;CPU Usage;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559920] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559935] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559938] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559950] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Zombie Process;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559953] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Zombie Process;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1456559953] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Total Processes;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1456559953] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Memory;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1456559953] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Root Partition;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1456559966] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Users;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1456559967] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Ping Check;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559967] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Ping Check;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1456559980] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1456559982] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Root Partition;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1456559998] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Total Processes;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1456560027] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Ping Check;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1456560042] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Root Partition;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1456560058] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Total Processes;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1456563750] SERVICE ALERT: credr_imageprocessing_52.76.65.82;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456563810] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1456563821] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Root Partition;OK;HARD;3;DISK OK - free space: / 10365 MB (54% inode=87%):
[1456563822] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Load;OK;HARD;3;OK - load average: 0.00, 0.01, 0.05
[1456563836] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Total Processes;OK;HARD;3;PROCS OK: 120 processes
[1456563894] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Current Users;OK;HARD;3;USERS OK - 1 users currently logged in
[1456563897] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Memory;OK;HARD;3;[MEMORY] Total: 3952 MB - Used: 269 MB - 6% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1456563909] SERVICE ALERT: credr_imageprocessing_52.76.65.82;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z