[1455569081] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=34.52% system=21.32%, iowait=38.58%, idle=5.58%, nice=0.00%, steal=0.00%
[1455569141] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=29.59% system=25.00%, iowait=39.29%, idle=5.61%, nice=0.00%, steal=0.51%
[1455569201] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=26.13% system=24.12%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455580869] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 132
[1455580929] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 126
[1455580989] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 258
[1455581169] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 119
[1455581350] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=28
[1455584501] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.00% system=4.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455584561] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=42.21% system=7.54%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455584621] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=46.00% system=3.50%, iowait=0.50%, idle=50.00%, nice=0.00%, steal=0.00%
[1455584801] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=31.50% system=1.00%, iowait=0.00%, idle=67.50%, nice=0.00%, steal=0.00%
[1455588041] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.50% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1455588101] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.28% system=5.47%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1455588161] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.78% system=5.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1455588701] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455591761] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.52% system=6.06%, iowait=0.00%, idle=42.42%, nice=0.00%, steal=0.00%
[1455591821] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=83.08% system=12.31%, iowait=0.00%, idle=4.62%, nice=0.00%, steal=0.00%
[1455591881] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=91.92% system=7.58%, iowait=0.00%, idle=0.51%, nice=0.00%, steal=0.00%
[1455592601] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455602141] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=16.08% system=4.02%, iowait=39.20%, idle=40.20%, nice=0.00%, steal=0.50%
[1455602202] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=3.03% system=3.03%, iowait=73.74%, idle=19.70%, nice=0.00%, steal=0.51%
[1455602261] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.50% system=0.50%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.00%
[1455603341] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.26% system=0.50%, iowait=0.00%, idle=48.24%, nice=0.00%, steal=0.00%
[1455603401] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455612581] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=2.05% system=2.05%, iowait=45.13%, idle=50.77%, nice=0.00%, steal=0.00%
[1455612641] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.01%, nice=0.00%, steal=0.50%
[1455613721] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.50% system=7.00%, iowait=0.00%, idle=48.50%, nice=0.00%, steal=0.00%
[1455613781] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.50% system=0.50%, iowait=0.00%, idle=98.50%, nice=0.00%, steal=0.50%
[1455614441] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=55.50% system=11.00%, iowait=0.50%, idle=32.50%, nice=0.00%, steal=0.50%
[1455614501] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455615761] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.00% system=7.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455615821] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%