[1453920814] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1453920874] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1453920934] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1453922194] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;OK;HARD;3;USERS OK - 4 users currently logged in
[1453922254] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;HARD;3;PROCS OK: 196 processes
[1453922434] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 203 processes
[1453922494] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 205 processes
[1453922554] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 200 processes
[1453922734] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 207 processes
[1453922794] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 203 processes
[1453922854] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;HARD;3;PROCS WARNING: 202 processes
[1453923754] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;HARD;3;PROCS OK: 200 processes
[1453923934] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 207 processes
[1453923994] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 202 processes
[1453924054] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 195 processes
[1453924234] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 202 processes
[1453924294] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 207 processes
[1453924354] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 196 processes
[1453924534] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 203 processes
[1453924594] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 202 processes
[1453924654] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 200 processes
[1453924834] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 202 processes
[1453924894] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 206 processes
[1453924954] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 195 processes
[1453925134] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 202 processes
[1453925194] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 207 processes
[1453925254] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 195 processes
[1453925434] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 203 processes
[1453925494] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 202 processes
[1453925554] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 200 processes
[1453925734] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 207 processes
[1453925794] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 203 processes
[1453925854] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 195 processes
[1453926034] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 202 processes
[1453926094] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 207 processes
[1453926154] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;3;PROCS OK: 195 processes
[1453926334] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 206 processes
[1453926394] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;2;PROCS OK: 199 processes
[1453939264] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 103
[1453939324] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;2;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 217
[1453939384] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 226
[1453939745] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1453956394] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 500 Server closed connection without sending any data back
[1453956394] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1453956454] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;SOFT;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1453965276] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=55.96% system=3.11%, iowait=0.00%, idle=40.41%, nice=0.00%, steal=0.52%
[1453965306] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=40.36% system=1.54%, iowait=0.00%, idle=57.84%, nice=0.00%, steal=0.26%
[1453965336] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=62.84% system=0.55%, iowait=7.65%, idle=28.96%, nice=0.00%, steal=0.00%
[1453965366] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=13.52% system=0.51%, iowait=0.00%, idle=85.97%, nice=0.00%, steal=0.00%
[1453965396] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=59.28% system=0.00%, iowait=0.00%, idle=40.72%, nice=0.00%, steal=0.00%
[1453965576] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=1.58% system=0.53%, iowait=0.00%, idle=96.84%, nice=0.00%, steal=1.05%
[1453965756] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=54.77% system=0.50%, iowait=0.00%, idle=44.72%, nice=0.00%, steal=0.00%
[1453965816] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=52.76% system=1.51%, iowait=0.00%, idle=45.73%, nice=0.00%, steal=0.00%
[1453965876] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=54.79% system=0.53%, iowait=0.00%, idle=44.68%, nice=0.00%, steal=0.00%
[1453966014] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=74.49% system=2.55%, iowait=0.00%, idle=22.45%, nice=0.00%, steal=0.51%
[1453966014] SERVICE FLAPPING ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1453966075] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.50%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.00%
[1453966260] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=98.01% system=1.99%, iowait=0.00%, idle=0.00%, nice=0.00%, steal=0.00%
[1453966320] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=63.49% system=0.00%, iowait=0.00%, idle=36.51%, nice=0.00%, steal=0.00%
[1453966380] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=56.19% system=2.06%, iowait=11.34%, idle=30.41%, nice=0.00%, steal=0.00%
[1453966560] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=14.21% system=0.51%, iowait=0.00%, idle=84.77%, nice=0.00%, steal=0.51%
[1453966894] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 205 processes
[1453966954] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;2;PROCS OK: 198 processes
[1453967134] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 206 processes
[1453967194] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 207 processes
[1453967254] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;HARD;3;PROCS WARNING: 205 processes
[1453968154] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;HARD;3;PROCS OK: 200 processes
[1453969980] SERVICE FLAPPING ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1453975534] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1453975594] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1453975654] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;SOFT;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1453976166] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=40.80% system=0.50%, iowait=0.00%, idle=58.71%, nice=0.00%, steal=0.00%
[1453976226] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=21.13% system=0.26%, iowait=0.00%, idle=78.61%, nice=0.00%, steal=0.00%
[1453991914] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 500 Server closed connection without sending any data back
[1453991973] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 500 Server closed connection without sending any data back
[1453992034] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1453992154] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;HARD;3;CONTENT OK: WEBSITE IS RUNNING FINE