[1451453054] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.02% system=8.38%, iowait=0.00%, idle=47.49%, nice=0.00%, steal=1.12%
[1451453174] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1451454383] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451454503] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451454623] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451457017] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451457036] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451457156] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451457276] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451458292] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451458292] SERVICE FLAPPING ALERT: credr_dev_52.74.244.220;HTTP;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1451458681] SERVICE FLAPPING ALERT: credr_dev_52.74.244.220;HTTP;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1451458723] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451458843] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451458963] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451459512] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451461850] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=60.80% system=3.02%, iowait=4.02%, idle=32.16%, nice=0.00%, steal=0.00%
[1451461970] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.51% system=0.51%, iowait=0.00%, idle=98.99%, nice=0.00%, steal=0.00%
[1451462893] SERVICE FLAPPING ALERT: credr_dev_52.74.244.220;HTTP;STARTED; Service appears to have started flapping (30.7% change >= 20.0% threshold)
[1451462951] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1451463071] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1451463141] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;UNKNOWN;HARD;3;NRPE: Unable to read output
[1451466959] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 203 processes
[1451467079] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;2;PROCS WARNING: 207 processes
[1451467199] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;HARD;3;PROCS WARNING: 206 processes
[1451467369] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1451468851] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451468957] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451469074] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451469194] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451469235] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451469258] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451469258] SERVICE FLAPPING ALERT: credr_dev_52.74.244.220;HTTP;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1451470458] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451470578] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451473799] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;HARD;3;PROCS OK: 199 processes
[1451474569] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=48.00% system=2.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1451474689] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1451476429] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476464] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476549] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476572] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476572] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476584] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476639] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476647] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476669] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476692] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476692] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476704] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476708] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476710] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476759] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476767] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476801] SERVICE ALERT: credr_stage_52.74.191.24;Memory;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1451476812] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476812] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476828] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476830] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476879] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476887] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476921] SERVICE ALERT: credr_stage_52.74.191.24;Memory;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1451476948] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451476950] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1451477041] SERVICE ALERT: credr_stage_52.74.191.24;Memory;CRITICAL;HARD;3;(Return code of 255 is out of bounds)