[1456425031] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.02% system=0.00%, iowait=1.02%, idle=47.45%, nice=0.00%, steal=0.51%
[1456425091] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=61.08% system=2.16%, iowait=0.00%, idle=36.76%, nice=0.00%, steal=0.00%
[1456425151] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.76% system=0.51%, iowait=0.00%, idle=48.22%, nice=0.00%, steal=0.51%
[1456425151] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=64.67% system=1.63%, iowait=0.54%, idle=32.61%, nice=0.00%, steal=0.54%
[1456425211] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=49.75% system=1.02%, iowait=0.00%, idle=48.73%, nice=0.00%, steal=0.51%
[1456425331] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=5.24% system=1.57%, iowait=0.00%, idle=92.15%, nice=0.00%, steal=1.05%
[1456427881] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=1.51%, iowait=0.00%, idle=48.24%, nice=0.00%, steal=0.00%
[1456427941] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456430101] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.32% system=2.46%, iowait=1.97%, idle=50.25%, nice=0.00%, steal=0.00%
[1456430161] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=4.02% system=0.00%, iowait=1.01%, idle=94.97%, nice=0.00%, steal=0.00%
[1456444950] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 100
[1456445009] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1456445069] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 231
[1456445249] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=93
[1456478161] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=4.25% system=95.75%, iowait=0.00%, idle=0.00%, nice=0.00%, steal=0.00%
[1456478221] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.25% system=0.00%, iowait=0.00%, idle=99.75%, nice=0.00%, steal=0.00%
[1456502071] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.01% system=1.01%, iowait=1.01%, idle=46.97%, nice=0.00%, steal=0.00%
[1456502131] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.50%, idle=98.99%, nice=0.00%, steal=0.00%