[1454783427] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=54.77% system=1.51%, iowait=4.02%, idle=39.20%, nice=0.00%, steal=0.50%
[1454783466] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=1.53%, iowait=66.84%, idle=31.12%, nice=0.00%, steal=0.51%
[1454783487] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=96.00% system=4.00%, iowait=0.00%, idle=0.00%, nice=0.00%, steal=0.00%
[1454783526] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=48.98% system=2.04%, iowait=0.00%, idle=48.47%, nice=0.00%, steal=0.51%
[1454783548] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=16.50% system=2.50%, iowait=46.50%, idle=34.00%, nice=0.00%, steal=0.50%
[1454783727] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454793423] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.00% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1454793483] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.49% system=0.00%, iowait=0.50%, idle=98.02%, nice=0.00%, steal=0.00%
[1454794023] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.00% system=6.50%, iowait=0.00%, idle=52.50%, nice=0.00%, steal=0.00%
[1454794083] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=6.00% system=0.00%, iowait=0.00%, idle=94.00%, nice=0.00%, steal=0.00%
[1454796423] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.56% system=6.44%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454796483] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454796543] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454796723] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=33.83% system=3.98%, iowait=0.00%, idle=62.19%, nice=0.00%, steal=0.00%
[1454803258] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454803319] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 140
[1454803378] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 217
[1454803558] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=99
[1454806803] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=7.33% system=2.62%, iowait=38.22%, idle=51.83%, nice=0.00%, steal=0.00%
[1454806863] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%