[1455733871] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.52% system=1.57%, iowait=60.21%, idle=36.65%, nice=0.00%, steal=1.05%
[1455733931] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.25% system=1.02%, iowait=0.51%, idle=47.72%, nice=0.00%, steal=0.51%
[1455733991] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=49.25% system=1.51%, iowait=0.00%, idle=48.74%, nice=0.00%, steal=0.50%
[1455753549] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 262
[1455753609] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 135
[1455753669] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 135
[1455754209] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=28
[1455757301] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455757361] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.27% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455757421] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=46.77% system=3.48%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455757601] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=31.19% system=0.99%, iowait=0.00%, idle=67.33%, nice=0.00%, steal=0.50%
[1455757751] SERVICE ALERT: Nagios Server;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=56.22% system=2.49%, iowait=8.46%, idle=32.34%, nice=0.00%, steal=0.50%
[1455757811] SERVICE ALERT: Nagios Server;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455760841] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.22% system=5.53%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455760901] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=47.00% system=7.00%, iowait=0.00%, idle=46.00%, nice=0.00%, steal=0.00%
[1455760961] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.05% system=4.46%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1455761681] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455764441] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.00% system=0.50%, iowait=0.00%, idle=48.00%, nice=0.00%, steal=0.50%
[1455764501] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.00% system=0.50%, iowait=0.00%, idle=98.51%, nice=0.00%, steal=0.00%
[1455764741] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455764801] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=38.50% system=11.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455764981] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.22% system=5.53%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455765041] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=43.50% system=8.50%, iowait=0.00%, idle=48.00%, nice=0.00%, steal=0.00%
[1455765101] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.72% system=6.03%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455765281] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=29.15% system=0.50%, iowait=0.00%, idle=69.85%, nice=0.00%, steal=0.50%
[1455765401] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.01% system=0.00%, iowait=0.00%, idle=48.99%, nice=0.00%, steal=0.00%
[1455765461] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.00% system=0.00%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.00%
[1455766541] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.75% system=0.50%, iowait=0.00%, idle=48.74%, nice=0.00%, steal=0.00%
[1455766601] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455775061] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=0.50%, iowait=0.00%, idle=48.74%, nice=0.00%, steal=0.50%
[1455775121] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.00% system=0.00%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.00%
[1455779559] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;SOFT;1;CRITICAL -  Error: Redis connection to localhost:6379 failed - connect ECONNREFUSED
[1455779560] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;SOFT;1;CRITICAL -  Error: Redis connection to localhost:6379 failed - connect ECONNREFUSED
[1455779619] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;OK;SOFT;2;OK - uptime_in_days: 0
[1455779619] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;OK;SOFT;2;OK - 511.13kB of memory used.
[1455779801] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.50% system=0.50%, iowait=0.00%, idle=48.00%, nice=0.00%, steal=0.00%
[1455779861] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455783821] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=81.82% system=17.17%, iowait=0.00%, idle=1.01%, nice=0.00%, steal=0.00%
[1455783881] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=74.37% system=24.62%, iowait=0.00%, idle=1.01%, nice=0.00%, steal=0.00%
[1455783941] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=76.00% system=22.50%, iowait=0.00%, idle=1.50%, nice=0.00%, steal=0.00%
[1455786101] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455788801] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.26% system=0.00%, iowait=0.00%, idle=48.74%, nice=0.00%, steal=0.00%
[1455788861] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=2.49% system=0.50%, iowait=0.00%, idle=97.01%, nice=0.00%, steal=0.00%
[1455790661] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.52% system=0.51%, iowait=0.00%, idle=47.98%, nice=0.00%, steal=0.00%
[1455790721] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.00% system=0.50%, iowait=0.00%, idle=96.50%, nice=0.00%, steal=0.00%
[1455793061] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=48.69% system=6.81%, iowait=0.00%, idle=43.46%, nice=0.00%, steal=1.05%
[1455793121] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.51% system=1.01%, iowait=0.00%, idle=48.48%, nice=0.00%, steal=0.00%
[1455793181] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.51% system=0.00%, iowait=0.00%, idle=99.49%, nice=0.00%, steal=0.00%
[1455793541] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.75% system=0.50%, iowait=0.00%, idle=48.74%, nice=0.00%, steal=0.00%
[1455793601] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.51% system=0.50%, iowait=0.00%, idle=97.99%, nice=0.00%, steal=0.00%
[1455801731] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=63.52% system=2.42%, iowait=0.00%, idle=33.93%, nice=0.00%, steal=0.13%
[1455801791] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=24.00% system=1.00%, iowait=0.25%, idle=74.62%, nice=0.00%, steal=0.12%
[1455801791] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=40.40% system=2.02%, iowait=2.02%, idle=55.56%, nice=0.00%, steal=0.00%
[1455801851] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455802241] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=69.23% system=4.10%, iowait=0.00%, idle=26.15%, nice=0.00%, steal=0.51%
[1455802301] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=68.39% system=4.66%, iowait=0.00%, idle=26.94%, nice=0.00%, steal=0.00%
[1455802361] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=74.00% system=6.50%, iowait=0.00%, idle=19.00%, nice=0.00%, steal=0.50%
[1455802901] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=1.00% system=0.00%, iowait=0.00%, idle=98.51%, nice=0.00%, steal=0.50%
[1455803441] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=69.39% system=3.57%, iowait=0.00%, idle=26.53%, nice=0.00%, steal=0.51%
[1455803501] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=69.23% system=3.08%, iowait=0.00%, idle=27.18%, nice=0.00%, steal=0.51%
[1455803561] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=70.77% system=3.59%, iowait=0.00%, idle=25.13%, nice=0.00%, steal=0.51%
[1455804322] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=69.43% system=4.15%, iowait=0.00%, idle=25.91%, nice=0.00%, steal=0.52%
[1455804322] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;CPU Usage;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1455804502] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=85.07% system=3.98%, iowait=0.00%, idle=10.95%, nice=0.00%, steal=0.00%
[1455804682] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=33.33% system=2.02%, iowait=0.00%, idle=64.65%, nice=0.00%, steal=0.00%
[1455804862] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=68.56% system=5.67%, iowait=0.00%, idle=25.26%, nice=0.00%, steal=0.52%
[1455804922] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=69.59% system=4.12%, iowait=0.00%, idle=25.77%, nice=0.00%, steal=0.52%
[1455804982] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=96.00% system=4.00%, iowait=0.00%, idle=0.00%, nice=0.00%, steal=0.00%
[1455805359] SERVICE ALERT: credr_stage_52.76.16.57;Nginx http;CRITICAL;SOFT;1;connect to address 52.76.16.57 and port 80: Connection refused
[1455805359] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455805359] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credrWebApp;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455805359] SERVICE ALERT: credr_stage_52.76.16.57;Memory;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455805359] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455805359] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455866548] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;CPU Usage;STARTED; Service appears to have started flapping (36.8% change >= 20.0% threshold)
[1455866554] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;OK;SOFT;2;OK - uptime_in_days: 0
[1455866574] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;OK;SOFT;2;Memory used by celery process :: 754.121 MB
[1455866586] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;OK;SOFT;2;OK - 511.13kB of memory used.
[1455866599] SERVICE ALERT: credr_stage_52.76.16.57;Nginx http;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 1099 bytes in 0.001 second response time
[1455866599] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credrWebApp;OK;SOFT;2;credrWebApp worker process running: 3
[1455866606] SERVICE ALERT: credr_stage_52.76.16.57;Memory;OK;SOFT;2;[MEMORY] Total: 16047 MB - Used: 2816 MB - 17% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1455866624] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455869684] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;CPU Usage;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1455874803] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor WebApp;WARNING;SOFT;1;credrWebApp                      STOPPING
[1455874863] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor WebApp;OK;SOFT;2;credrWebApp                      RUNNING    pid 32188, uptime 0:00:56
[1455894023] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9997 MB (26% inode=97%):
[1455894083] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9990 MB (25% inode=97%):
[1455894143] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;HARD;3;DISK WARNING - free space: / 9997 MB (26% inode=97%):