[1454007934] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;1;PROCS WARNING: 298 processes
[1454007994] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;2;PROCS WARNING: 294 processes
[1454008054] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;HARD;3;PROCS WARNING: 290 processes
[1454009494] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;OK;HARD;3;PROCS OK: 245 processes
[1454025784] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 217
[1454025844] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;2;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 226
[1454025904] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 138
[1454026085] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1454038474] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;1;PROCS WARNING: 274 processes
[1454038534] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;2;PROCS WARNING: 262 processes
[1454038594] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;HARD;3;PROCS WARNING: 260 processes
[1454038774] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;OK;HARD;3;PROCS OK: 243 processes
[1454040126] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=4.02%, iowait=0.00%, idle=45.73%, nice=0.00%, steal=0.00%
[1454040186] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=41.00% system=3.50%, iowait=0.00%, idle=55.50%, nice=0.00%, steal=0.00%
[1454050206] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=63.28% system=1.74%, iowait=0.00%, idle=34.99%, nice=0.00%, steal=0.00%
[1454050266] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=13.55% system=0.51%, iowait=0.00%, idle=85.93%, nice=0.00%, steal=0.00%
[1454051374] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;1;PROCS WARNING: 263 processes
[1454051434] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;OK;SOFT;2;PROCS OK: 234 processes
[1454052951] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=56.15% system=0.53%, iowait=0.00%, idle=43.32%, nice=0.00%, steal=0.00%
[1454053011] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=13.64% system=0.51%, iowait=0.00%, idle=85.86%, nice=0.00%, steal=0.00%
[1454060706] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=63.54% system=2.03%, iowait=0.00%, idle=34.18%, nice=0.00%, steal=0.25%
[1454060766] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=15.37% system=0.50%, iowait=0.00%, idle=84.13%, nice=0.00%, steal=0.00%
[1454062684] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9998 MB (26% inode=97%):
[1454062744] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9998 MB (26% inode=97%):
[1454062804] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;3;DISK OK - free space: / 10004 MB (26% inode=97%):
[1454063164] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9994 MB (25% inode=97%):
[1454063224] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10007 MB (26% inode=97%):
[1454063404] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9998 MB (26% inode=97%):
[1454063464] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10005 MB (26% inode=97%):
[1454063644] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9989 MB (25% inode=97%):
[1454063704] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9993 MB (25% inode=97%):
[1454063764] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;3;DISK OK - free space: / 10001 MB (26% inode=97%):
[1454063944] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9999 MB (26% inode=97%):
[1454064004] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9994 MB (25% inode=97%):
[1454064064] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;HARD;3;DISK WARNING - free space: / 9983 MB (25% inode=97%):
[1454064906] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.69% system=1.52%, iowait=0.00%, idle=54.80%, nice=0.00%, steal=0.00%
[1454064966] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=19.47% system=0.00%, iowait=0.00%, idle=80.53%, nice=0.00%, steal=0.00%
[1454066404] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;HARD;3;DISK OK - free space: / 10315 MB (26% inode=97%):
[1454067306] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.54% system=1.52%, iowait=0.00%, idle=54.94%, nice=0.00%, steal=0.00%
[1454067366] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=25.71% system=1.04%, iowait=0.26%, idle=72.73%, nice=0.00%, steal=0.26%
[1454068806] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.46% system=2.01%, iowait=0.00%, idle=56.53%, nice=0.00%, steal=0.00%
[1454068866] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=26.50% system=1.00%, iowait=0.00%, idle=72.50%, nice=0.00%, steal=0.00%
[1454069046] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.70% system=2.28%, iowait=0.00%, idle=50.76%, nice=0.00%, steal=0.25%
[1454069106] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=18.94% system=0.25%, iowait=0.00%, idle=80.56%, nice=0.00%, steal=0.25%
[1454070514] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 360 processes
[1454070546] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=48.32% system=3.36%, iowait=0.26%, idle=48.06%, nice=0.00%, steal=0.00%
[1454070574] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 372 processes
[1454070606] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=27.44% system=0.77%, iowait=0.00%, idle=71.79%, nice=0.00%, steal=0.00%
[1454070634] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 374 processes
[1454071534] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;HARD;3;PROCS WARNING: 342 processes
[1454071894] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;OK;HARD;3;PROCS OK: 171 processes
[1454084854] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;1;PROCS WARNING: 263 processes
[1454084914] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;SOFT;2;PROCS WARNING: 305 processes
[1454084974] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;WARNING;HARD;3;PROCS WARNING: 291 processes
[1454085154] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;Total Processes;OK;HARD;3;PROCS OK: 173 processes