[1456252231] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=1.05% system=3.14%, iowait=62.83%, idle=31.94%, nice=0.00%, steal=1.05%
[1456252291] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.51% system=1.02%, iowait=81.12%, idle=16.84%, nice=0.00%, steal=0.51%
[1456252351] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=49.75% system=1.51%, iowait=0.50%, idle=47.74%, nice=0.00%, steal=0.50%
[1456271969] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 115
[1456272030] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 135
[1456272089] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 127
[1456272269] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 258
[1456272449] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 120
[1456272630] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=66
[1456272630] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;PG Locks;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1456275635] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=1.00%, idle=48.50%, nice=0.00%, steal=0.50%
[1456275695] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.00% system=6.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456275755] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=42.08% system=7.92%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1456276049] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;PG Locks;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1456276295] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1456279355] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.00% system=6.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456279415] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456279475] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=46.50% system=3.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456280195] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456282895] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.27% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1456282955] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=43.78% system=6.47%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1456283015] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.77% system=3.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1456283735] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456290395] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=14.14% system=3.54%, iowait=50.51%, idle=31.31%, nice=0.00%, steal=0.51%
[1456290455] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=8.54% system=2.01%, iowait=74.87%, idle=14.07%, nice=0.00%, steal=0.50%
[1456290515] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456299335] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=7.18% system=1.54%, iowait=40.51%, idle=50.77%, nice=0.00%, steal=0.00%
[1456299395] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.00% system=0.51%, iowait=48.98%, idle=50.51%, nice=0.00%, steal=0.00%
[1456299456] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=4.04% system=1.01%, iowait=51.52%, idle=43.43%, nice=0.00%, steal=0.00%
[1456300175] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456300355] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=9.50% system=1.00%, iowait=39.50%, idle=50.00%, nice=0.00%, steal=0.00%
[1456300416] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.51% system=0.00%, iowait=53.85%, idle=45.64%, nice=0.00%, steal=0.00%
[1456300475] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=2.55% system=2.04%, iowait=47.45%, idle=47.96%, nice=0.00%, steal=0.00%
[1456300655] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456300655] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1456302635] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=3.55% system=2.03%, iowait=44.16%, idle=50.25%, nice=0.00%, steal=0.00%
[1456302695] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=2.04% system=2.04%, iowait=39.29%, idle=56.63%, nice=0.00%, steal=0.00%
[1456302755] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=2.03% system=1.52%, iowait=45.69%, idle=50.76%, nice=0.00%, steal=0.00%
[1456303115] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=33.17% system=3.52%, iowait=39.20%, idle=24.12%, nice=0.00%, steal=0.00%
[1456303171] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=52.74% system=1.49%, iowait=0.00%, idle=44.78%, nice=0.00%, steal=1.00%
[1456303231] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=63.59% system=5.13%, iowait=0.51%, idle=30.26%, nice=0.00%, steal=0.51%
[1456303291] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=35.75% system=8.29%, iowait=21.24%, idle=33.68%, nice=0.00%, steal=1.04%
[1456303295] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=2.59% system=4.66%, iowait=40.93%, idle=51.30%, nice=0.00%, steal=0.52%
[1456303655] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=0.00% system=0.00%, iowait=37.50%, idle=61.98%, nice=0.00%, steal=0.52%
[1456303835] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=1.55% system=1.03%, iowait=51.55%, idle=45.88%, nice=0.00%, steal=0.00%
[1456304195] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=1.50% system=1.00%, iowait=0.50%, idle=97.00%, nice=0.00%, steal=0.00%
[1456306535] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=3.03% system=2.02%, iowait=44.44%, idle=50.51%, nice=0.00%, steal=0.00%
[1456306595] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=5.64% system=1.54%, iowait=42.05%, idle=50.77%, nice=0.00%, steal=0.00%
[1456306655] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=11.94% system=1.99%, iowait=36.82%, idle=49.25%, nice=0.00%, steal=0.00%
[1456306835] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=16.84% system=2.55%, iowait=29.59%, idle=51.02%, nice=0.00%, steal=0.00%
[1456309534] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=1.03% system=1.55%, iowait=46.39%, idle=50.52%, nice=0.00%, steal=0.52%
[1456309594] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.00% system=0.50%, iowait=0.50%, idle=98.01%, nice=0.00%, steal=0.00%
[1456310314] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1456338479] SERVICE ALERT: credr_stage_52.76.16.57;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 500 Can't connect to credr.com:80 (Connection refused)
[1456338539] SERVICE ALERT: credr_stage_52.76.16.57;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 500 Can't connect to credr.com:80 (Connection refused)
[1456338539] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 500 Can't connect to credr.com:80 (Connection refused)
[1456338569] SERVICE ALERT: credr_prod_52.74.247.15;Nginx http;CRITICAL;SOFT;1;connect to address 52.74.247.15 and port 80: Connection refused
[1456338599] SERVICE ALERT: credr_stage_52.76.16.57;Website Check;UNKNOWN;HARD;3;ERROR: CANNOT RETRIEVE URL: 500 Can't connect to credr.com:80 (Connection refused)
[1456338599] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 500 Can't connect to credr.com:80 (Connection refused)