[1455042689] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=2.62%, iowait=54.97%, idle=41.88%, nice=0.00%, steal=0.52%
[1455042709] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=60.56% system=1.67%, iowait=0.56%, idle=37.22%, nice=0.00%, steal=0.00%
[1455042749] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=48.99% system=2.53%, iowait=0.00%, idle=47.98%, nice=0.00%, steal=0.51%
[1455042768] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=20.33% system=1.65%, iowait=0.00%, idle=77.47%, nice=0.00%, steal=0.55%
[1455045823] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=2.49%, iowait=0.00%, idle=47.26%, nice=0.00%, steal=0.00%
[1455045883] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455062369] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 230
[1455062428] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 181
[1455062488] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1455062668] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 226
[1455063029] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1455075641] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.23% system=5.03%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455075701] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=25.63% system=5.53%, iowait=0.00%, idle=68.84%, nice=0.00%, steal=0.00%
[1455076421] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=47.76% system=1.99%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1455076481] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455076541] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.23% system=4.52%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455076721] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=28.71% system=2.48%, iowait=0.00%, idle=68.81%, nice=0.00%, steal=0.00%
[1455077081] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=5.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1455077141] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=4.50% system=0.00%, iowait=0.00%, idle=95.00%, nice=0.00%, steal=0.50%
[1455097645] SERVICE ALERT: credr_stage_52.76.16.57;Current Load;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097659] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credrWebApp;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097662] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097664] SERVICE ALERT: credr_stage_52.76.16.57;Nginx http;CRITICAL;SOFT;1;connect to address 52.76.16.57 and port 80: Connection refused
[1455097674] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097678] SERVICE ALERT: credr_stage_52.76.16.57;Memory;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097689] SERVICE ALERT: credr_stage_52.76.16.57;Ping Check;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097694] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process crmserv;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097695] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097705] SERVICE ALERT: credr_stage_52.76.16.57;Current Load;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097710] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process Total;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097713] SERVICE ALERT: credr_stage_52.76.16.57;Root Partition;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097719] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory emailserv;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097719] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097719] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credrWebApp;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097719] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credr-core;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097722] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097724] SERVICE ALERT: credr_stage_52.76.16.57;SSH;CRITICAL;SOFT;1;connect to address 52.76.16.57 and port 22: Connection refused
[1455097724] SERVICE ALERT: credr_stage_52.76.16.57;Nginx http;CRITICAL;SOFT;2;connect to address 52.76.16.57 and port 80: Connection refused
[1455097725] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097734] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097738] SERVICE ALERT: credr_stage_52.76.16.57;Memory;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097749] SERVICE ALERT: credr_stage_52.76.16.57;Ping Check;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097752] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credrWebApp;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097754] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process crmserv;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097755] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097765] SERVICE ALERT: credr_stage_52.76.16.57;Current Load;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097769] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097770] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process Total;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097772] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory crmserv;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097772] SERVICE ALERT: credr_stage_52.76.16.57;Root Partition;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097776] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory smsserv;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097779] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory emailserv;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097779] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097779] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credrWebApp;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097779] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credr-core;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097781] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process emailserv;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097782] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097784] SERVICE ALERT: credr_stage_52.76.16.57;SSH;CRITICAL;SOFT;2;connect to address 52.76.16.57 and port 22: Connection refused
[1455097784] SERVICE ALERT: credr_stage_52.76.16.57;Nginx http;CRITICAL;HARD;3;connect to address 52.76.16.57 and port 80: Connection refused
[1455097785] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097793] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097794] SERVICE ALERT: credr_stage_52.76.16.57;Total Processes;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097798] SERVICE ALERT: credr_stage_52.76.16.57;Memory;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097809] SERVICE ALERT: credr_stage_52.76.16.57;Zombie Process;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097809] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process smsserv;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455097809] SERVICE ALERT: credr_stage_52.76.16.57;Ping Check;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097812] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credrWebApp;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097814] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process crmserv;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097815] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097829] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097830] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process Total;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097832] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory crmserv;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097832] SERVICE ALERT: credr_stage_52.76.16.57;Root Partition;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097836] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory smsserv;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097839] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097839] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory emailserv;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097839] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credr-core;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097841] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process emailserv;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097844] SERVICE ALERT: credr_stage_52.76.16.57;SSH;CRITICAL;HARD;3;connect to address 52.76.16.57 and port 22: Connection refused
[1455097845] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097854] SERVICE ALERT: credr_stage_52.76.16.57;Total Processes;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097869] SERVICE ALERT: credr_stage_52.76.16.57;Zombie Process;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097869] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process smsserv;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1455097872] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credrWebApp;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097889] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097892] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory crmserv;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097896] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory smsserv;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097901] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process emailserv;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097914] SERVICE ALERT: credr_stage_52.76.16.57;Total Processes;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097929] SERVICE ALERT: credr_stage_52.76.16.57;Zombie Process;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455097929] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process smsserv;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1455098184] HOST ALERT: credr_stage_52.76.16.57;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455098207] SERVICE ALERT: credr_stage_52.76.16.57;Root Partition;OK;HARD;3;DISK OK - free space: / 25143 MB (87% inode=85%):
[1455098211] HOST ALERT: credr_stage_52.76.16.57;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.62 ms
[1455098214] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory emailserv;OK;HARD;3;Memory used by celery_emailserv process :: 35.0234 MB
[1455098214] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credr-core;OK;HARD;3;celery_credrcore worker process running: 1
[1455098221] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process Total;OK;HARD;3;celery worker process running: 4
[1455098232] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credrWebApp;OK;HARD;3;Memory used by credrWebApp process :: 67.6758 MB
[1455098249] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;OK;HARD;3;USERS OK - 0 users currently logged in
[1455098252] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory crmserv;OK;HARD;3;Memory used by celery_crmserv process :: 3.25 MB
[1455098256] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory smsserv;OK;HARD;3;Memory used by celery_smsserv process :: 3.26562 MB
[1455098261] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process emailserv;OK;HARD;3;emailserv worker process running: 1
[1455098274] SERVICE ALERT: credr_stage_52.76.16.57;Total Processes;OK;HARD;3;PROCS OK: 125 processes
[1455098289] SERVICE ALERT: credr_stage_52.76.16.57;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1455098289] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process smsserv;OK;HARD;3;smsserv worker process running: 1
[1455098305] SERVICE ALERT: credr_stage_52.76.16.57;Current Load;OK;HARD;3;OK - load average: 1.33, 0.61, 0.23
[1455098319] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credrWebApp;OK;HARD;3;credrWebApp worker process running: 3
[1455098322] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;OK;HARD;3;Memory used by celery process :: 3.24609 MB
[1455098324] SERVICE ALERT: credr_stage_52.76.16.57;Nginx http;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 1099 bytes in 0.003 second response time
[1455098333] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;OK;HARD;3;OK - uptime_in_days: 0
[1455098338] SERVICE ALERT: credr_stage_52.76.16.57;Memory;OK;HARD;3;[MEMORY] Total: 7984 MB - Used: 554 MB - 6% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1455098353] SERVICE ALERT: credr_stage_52.76.16.57;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.02 ms
[1455098354] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process crmserv;OK;HARD;3;crmserv worker process running: 3
[1455098355] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;OK;HARD;3;Memory used by celery_credrcore process :: 3.25 MB
[1455098381] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=39.70% system=11.06%, iowait=0.00%, idle=49.25%, nice=0.00%, steal=0.00%
[1455098384] SERVICE ALERT: credr_stage_52.76.16.57;SSH;OK;HARD;3;SSH OK - OpenSSH_6.6.1p1 Ubuntu-2ubuntu2.6 (protocol 2.0)
[1455098386] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;OK;HARD;3;OK - 489.75kB of memory used.
[1455102881] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.69% system=4.06%, iowait=0.51%, idle=49.24%, nice=0.00%, steal=0.51%
[1455102941] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.51% system=0.00%, iowait=0.00%, idle=98.49%, nice=0.00%, steal=0.00%
[1455104741] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=49.25% system=4.02%, iowait=0.00%, idle=46.73%, nice=0.00%, steal=0.00%
[1455104801] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.01% system=0.00%, iowait=0.00%, idle=98.99%, nice=0.00%, steal=0.00%
[1455115241] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=48.24% system=3.52%, iowait=0.00%, idle=47.74%, nice=0.00%, steal=0.50%
[1455115301] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.00% system=0.50%, iowait=0.00%, idle=98.51%, nice=0.00%, steal=0.00%