[1455388271] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=2.67%, iowait=49.73%, idle=47.06%, nice=0.00%, steal=0.53%
[1455388331] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=49.24% system=1.52%, iowait=0.51%, idle=47.72%, nice=0.00%, steal=1.02%
[1455399821] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.27% system=3.98%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.50%
[1455399881] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.00% system=6.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455399941] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=44.06% system=5.94%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1455401201] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455401381] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.22% system=5.53%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455401441] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=56.22% system=5.47%, iowait=0.00%, idle=38.31%, nice=0.00%, steal=0.00%
[1455401501] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.77% system=4.48%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455402221] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455402221] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1455405101] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1455408009] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 132
[1455408069] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 132
[1455408130] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 258
[1455408309] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=28
[1455427719] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;SOFT;1;CRITICAL -  Error: Redis connection to localhost:6379 failed - connect ECONNREFUSED
[1455427779] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;OK;SOFT;2;OK - uptime_in_days: 0
[1455432521] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=5.03% system=2.01%, iowait=42.71%, idle=49.75%, nice=0.00%, steal=0.50%
[1455432581] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.50%