[1452049777] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.23% system=4.02%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1452049837] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=36.00% system=6.00%, iowait=0.00%, idle=58.00%, nice=0.00%, steal=0.00%
[1452050017] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=87.74% system=8.39%, iowait=0.00%, idle=3.23%, nice=0.00%, steal=0.65%
[1452050077] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=77.03% system=10.81%, iowait=0.00%, idle=11.49%, nice=0.00%, steal=0.68%
[1452050137] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=69.73% system=16.22%, iowait=0.00%, idle=13.51%, nice=0.00%, steal=0.54%
[1452050317] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1452050497] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.05% system=5.45%, iowait=0.00%, idle=49.01%, nice=0.00%, steal=0.50%
[1452050557] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=73.91% system=13.59%, iowait=0.00%, idle=11.96%, nice=0.00%, steal=0.54%
[1452050617] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=86.90% system=8.97%, iowait=0.00%, idle=3.45%, nice=0.00%, steal=0.69%
[1452050797] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=29.65% system=6.53%, iowait=0.00%, idle=63.82%, nice=0.00%, steal=0.00%
[1452050797] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1452052056] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=47.26% system=2.99%, iowait=0.00%, idle=49.25%, nice=0.00%, steal=0.50%
[1452052116] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1452052176] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=25.63% system=3.52%, iowait=0.00%, idle=70.85%, nice=0.00%, steal=0.00%
[1452053076] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.73% system=4.02%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1452053136] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=40.00% system=4.00%, iowait=0.00%, idle=55.50%, nice=0.00%, steal=0.50%
[1452053196] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=33.17% system=4.52%, iowait=0.00%, idle=62.31%, nice=0.00%, steal=0.00%
[1452053736] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.78% system=6.47%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1452053796] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=22.39% system=3.48%, iowait=0.00%, idle=73.63%, nice=0.00%, steal=0.50%
[1452053976] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.71% system=5.03%, iowait=0.00%, idle=53.27%, nice=0.00%, steal=0.00%
[1452054036] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=46.50% system=3.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1452054096] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=36.68% system=2.51%, iowait=0.00%, idle=60.80%, nice=0.00%, steal=0.00%
[1452054636] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452063465] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9509 MB (24% inode=97%):
[1452063469] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9509 MB (24% inode=97%):
[1452063477] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;HARD;3;DISK WARNING - free space: / 9509 MB (24% inode=97%):
[1452063591] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1452063643] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;SOFT;2;CONTENT OK: WEBSITE IS RUNNING FINE
[1452063657] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;HARD;3;DISK OK - free space: / 12522 MB (32% inode=97%):
[1452064356] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.73% system=4.52%, iowait=0.00%, idle=48.74%, nice=0.00%, steal=0.00%
[1452064416] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.00% system=0.00%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.00%
[1452074854] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.78% system=5.47%, iowait=0.00%, idle=49.25%, nice=0.00%, steal=0.50%
[1452074914] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.50%, idle=99.50%, nice=0.00%, steal=0.00%
[1452075094] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=51.00% system=1.00%, iowait=0.50%, idle=47.50%, nice=0.00%, steal=0.00%
[1452075154] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.50%, idle=99.00%, nice=0.00%, steal=0.00%
[1452075334] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=47.47% system=0.00%, iowait=0.00%, idle=52.02%, nice=0.00%, steal=0.51%
[1452075394] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=47.50% system=2.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1452075454] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=46.73% system=3.02%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1452075546] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=46.73% system=3.52%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1452077194] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.71% system=0.49%, iowait=0.00%, idle=57.80%, nice=0.00%, steal=0.00%
[1452077254] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=12.28% system=1.00%, iowait=0.00%, idle=86.72%, nice=0.00%, steal=0.00%
[1452081325] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;CRITICAL;SOFT;1;(Service check timed out after 60.02 seconds)
[1452081325] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1452081367] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;HARD;3;ERROR: CANNOT RETRIEVE URL: 500 Server closed connection without sending any data back
[1452081614] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;HARD;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1452081925] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1452081985] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1452082045] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;CRITICAL;HARD;3;CONTENT ERROR: EXPR NOT FOUND, WEBSITE MIGHT BE DOWN !! (last: </html>)
[1452082489] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;HARD;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1452082489] SERVICE FLAPPING ALERT: credr_prod_52.74.247.15;Website Check;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1452082952] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor WebApp;WARNING;SOFT;1;credrWebApp                      STOPPING
[1452083012] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor WebApp;OK;SOFT;2;credrWebApp                      RUNNING    pid 31571, uptime 0:00:08
[1452085549] SERVICE FLAPPING ALERT: credr_prod_52.74.247.15;Website Check;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)