[1451284155] HOST ALERT: credr_prod_52.74.247.15;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1451284799] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;CRITICAL;SOFT;1;CRITICAL: Could not get status of credrWebApp
[1451284802] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;CRITICAL;SOFT;2;CRITICAL: Could not get status of credrWebApp
[1451284922] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1451285501] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285505] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285511] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285531] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451285536] SERVICE ALERT: credr_dev_52.74.244.220;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285551] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285556] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285566] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451285581] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451285618] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285648] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451285656] SERVICE ALERT: credr_dev_52.74.244.220;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285656] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451285658] SERVICE ALERT: credr_dev_52.74.244.220;Zombie Process;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285671] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285676] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285688] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451285701] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451285721] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285738] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285749] SERVICE ALERT: credr_dev_52.74.244.220;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285751] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451285776] SERVICE ALERT: credr_dev_52.74.244.220;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285778] SERVICE ALERT: credr_dev_52.74.244.220;Zombie Process;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285778] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451285785] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285791] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285796] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285815] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451285818] SERVICE ALERT: credr_dev_52.74.244.220;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285821] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451285841] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285858] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285869] SERVICE ALERT: credr_dev_52.74.244.220;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285871] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451285873] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285898] SERVICE ALERT: credr_dev_52.74.244.220;Zombie Process;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285899] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451285905] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285908] SERVICE ALERT: credr_stage_52.74.191.24;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285935] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451285938] SERVICE ALERT: credr_dev_52.74.244.220;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285961] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285968] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451285969] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285989] SERVICE ALERT: credr_dev_52.74.244.220;Current Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285989] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451285991] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451285993] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286018] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286019] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451286021] SERVICE ALERT: credr_dev_52.74.244.220;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286025] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286028] SERVICE ALERT: credr_stage_52.74.191.24;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286048] HOST ALERT: credr_stage_52.74.191.24;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1451286051] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451286058] SERVICE ALERT: credr_dev_52.74.244.220;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286088] HOST ALERT: credr_dev_52.74.244.220;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1451286089] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286109] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286113] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286138] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286141] SERVICE ALERT: credr_dev_52.74.244.220;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286148] SERVICE ALERT: credr_stage_52.74.191.24;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286209] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286229] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286258] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286261] SERVICE ALERT: credr_dev_52.74.244.220;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1451286863] SERVICE ALERT: credr_stage_52.74.191.24;Root Partition;OK;HARD;3;DISK OK - free space: / 51740 MB (89% inode=92%):
[1451286864] HOST ALERT: credr_dev_52.74.244.220;UP;HARD;10;PING OK - Packet loss = 68%, RTA = 0.67 ms
[1451286866] SERVICE ALERT: credr_dev_52.74.244.220;Total Processes;OK;HARD;3;PROCS OK: 122 processes
[1451286867] HOST ALERT: credr_stage_52.74.191.24;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1451286872] SERVICE ALERT: credr_stage_52.74.191.24;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1451286895] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;OK;HARD;3;credrWebApp                      RUNNING    pid 27979, uptime 1 day, 23:14:51
[1451286966] SERVICE ALERT: credr_dev_52.74.244.220;Memory;OK;HARD;3;[MEMORY] Total: 3951 MB - Used: 2598 MB - 65% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1451286986] SERVICE ALERT: credr_stage_52.74.191.24;Total Processes;OK;HARD;3;PROCS OK: 133 processes
[1451287048] SERVICE ALERT: credr_stage_52.74.191.24;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1451287061] SERVICE ALERT: credr_stage_52.74.191.24;Current Load;OK;HARD;3;OK - load average: 0.00, 0.01, 0.05
[1451287088] SERVICE ALERT: credr_dev_52.74.244.220;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1451287116] SERVICE ALERT: credr_stage_52.74.191.24;Current Users;OK;HARD;3;USERS OK - 1 users currently logged in
[1451287128] SERVICE ALERT: credr_stage_52.74.191.24;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1451287179] SERVICE ALERT: credr_dev_52.74.244.220;Current Load;OK;HARD;3;OK - load average: 0.00, 0.01, 0.05
[1451287194] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451287197] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451287206] SERVICE ALERT: credr_stage_52.74.191.24;Memory;OK;HARD;3;[MEMORY] Total: 3951 MB - Used: 498 MB - 12% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1451287223] SERVICE ALERT: credr_dev_52.74.244.220;Current Users;OK;HARD;3;USERS OK - 0 users currently logged in
[1451287236] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1451287236] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451287242] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;OK;HARD;3;DISK OK - free space: / 1183 MB (15% inode=62%):
[1451287253] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451287309] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451287349] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451287408] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451287439] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451287456] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451287481] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451287511] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451287571] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451287578] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451287616] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451287621] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451287678] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451287711] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451287746] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451287801] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451287836] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451287853] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451287933] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451287940] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451287944] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451287970] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451288023] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451288034] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451288056] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451288081] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451288111] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451288171] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451288178] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451288216] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451288221] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451288278] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451288311] HOST ALERT: credr_dev_52.74.244.220;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451288346] HOST ALERT: credr_stage_52.74.191.24;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451288401] HOST ALERT: credr_dev_52.74.244.220;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1451288436] HOST ALERT: credr_stage_52.74.191.24;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1451290426] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;SOFT;1;CRITICAL: Could not get status of credrWebApp
[1451290546] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;SOFT;2;CRITICAL: Could not get status of credrWebApp
[1451290666] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1451291677] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;OK;HARD;3;credrWebApp                      RUNNING    pid 27979, uptime 2 days, 0:34:34
[1451291677] SERVICE FLAPPING ALERT: credr_stage_52.74.191.24;Supervisor;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1451291685] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;OK;HARD;3;credrWebApp                      RUNNING    pid 13663, uptime 23:00:31
[1451293998] SERVICE FLAPPING ALERT: credr_stage_52.74.191.24;Supervisor;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1451296787] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451296812] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1451296842] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451296898] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451296932] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1451296954] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451297010] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451297036] SERVICE ALERT: credr_prodsolr_52.74.114.132;Nginx http;CRITICAL;SOFT;1;connect to address 52.74.114.132 and port 80: Connection refused
[1451297052] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;UNKNOWN;HARD;3;NRPE: Unable to read output
[1451297066] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451297101] SERVICE ALERT: credr_prodsolr_52.74.114.132;Ping Check;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451297131] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451297156] SERVICE ALERT: credr_prodsolr_52.74.114.132;Nginx http;CRITICAL;SOFT;2;connect to address 52.74.114.132 and port 80: Connection refused
[1451297177] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451297221] SERVICE ALERT: credr_prodsolr_52.74.114.132;Ping Check;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451297233] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451297275] SERVICE ALERT: credr_prodsolr_52.74.114.132;Nginx http;CRITICAL;HARD;3;connect to address 52.74.114.132 and port 80: Connection refused
[1451297289] HOST ALERT: credr_prodsolr_52.74.114.132;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1451297341] SERVICE ALERT: credr_prodsolr_52.74.114.132;Ping Check;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1451297718] SERVICE ALERT: credr_prodsolr_52.74.114.132;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1451298053] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1451299442] SERVICE ALERT: credr_prod_52.74.247.15;Ping Check;CRITICAL;SOFT;1;NRPE: Command 'check_ping' not defined
[1451299562] SERVICE ALERT: credr_prod_52.74.247.15;Ping Check;CRITICAL;SOFT;2;NRPE: Command 'check_ping' not defined
[1451299682] SERVICE ALERT: credr_prod_52.74.247.15;Ping Check;CRITICAL;HARD;3;NRPE: Command 'check_ping' not defined
[1451301478] SERVICE ALERT: credr_prod_52.74.247.15;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.02 ms
[1451301833] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1451301893] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 80: Connection refused
[1451301953] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 80: Connection refused
[1451302013] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 80: Connection refused
[1451302677] SERVICE ALERT: Nagios Server;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 11783 bytes in 0.001 second response time
[1451302977] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1451303037] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 80: Connection refused
[1451303097] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 80: Connection refused
[1451303157] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 80: Connection refused
[1451303757] SERVICE ALERT: Nagios Server;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 11783 bytes in 0.000 second response time
[1451303757] SERVICE FLAPPING ALERT: Nagios Server;HTTP;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1451304365] SERVICE FLAPPING ALERT: Nagios Server;HTTP;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1451304405] SERVICE FLAPPING ALERT: Nagios Server;HTTP;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1451304720] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1451304779] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 80: Connection refused
[1451304839] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 80: Connection refused
[1451304899] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 80: Connection refused
[1451305199] SERVICE ALERT: Nagios Server;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 11783 bytes in 0.001 second response time
[1451305499] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1451305559] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 80: Connection refused
[1451305619] SERVICE ALERT: Nagios Server;HTTP;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 11783 bytes in 0.001 second response time
[1451306519] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1451306579] SERVICE ALERT: Nagios Server;HTTP;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 80: Connection refused
[1451306639] SERVICE ALERT: Nagios Server;HTTP;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 11783 bytes in 0.001 second response time
[1451311139] SERVICE FLAPPING ALERT: Nagios Server;HTTP;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)