[1452227284] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=68.25% system=13.76%, iowait=0.00%, idle=16.93%, nice=0.00%, steal=1.06%
[1452227344] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=80.21% system=11.98%, iowait=0.00%, idle=2.60%, nice=0.00%, steal=5.21%
[1452227404] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=80.48% system=12.38%, iowait=0.00%, idle=0.95%, nice=0.00%, steal=6.19%
[1452227584] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=3.52% system=0.50%, iowait=0.00%, idle=95.98%, nice=0.00%, steal=0.00%
[1452227764] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=88.59% system=6.71%, iowait=0.00%, idle=3.36%, nice=0.00%, steal=1.34%
[1452227824] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=44.50% system=5.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1452229712] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes
[1452229772] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;2;PROCS OK: 245 processes
[1452232894] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=64.83% system=1.57%, iowait=0.00%, idle=33.60%, nice=0.00%, steal=0.00%
[1452232954] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.02% system=0.25%, iowait=0.00%, idle=96.73%, nice=0.00%, steal=0.00%
[1452234452] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes
[1452234512] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes
[1452234572] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;3;PROCS OK: 250 processes
[1452234752] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes
[1452234812] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;2;PROCS WARNING: 255 processes
[1452234872] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;HARD;3;PROCS WARNING: 253 processes
[1452235052] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;HARD;3;PROCS OK: 250 processes
[1452235474] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.27% system=1.80%, iowait=0.00%, idle=55.67%, nice=0.00%, steal=0.26%
[1452235534] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=16.49% system=0.52%, iowait=0.00%, idle=82.73%, nice=0.00%, steal=0.26%
[1452235772] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes
[1452235832] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;2;PROCS OK: 234 processes
[1452236637] SERVICE ALERT: credr_DW_52.76.17.124;PG DB Size;CRITICAL;SOFT;1;NRPE: Unable to read output
[1452236646] SERVICE ALERT: credr_DW_52.76.17.124;PG Ideal Connections;CRITICAL;SOFT;1;POSTGRES_TXN_IDLE CRITICAL: DB "apps_data" Total idle in transaction: 72
[1452236683] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 268
[1452236697] SERVICE ALERT: credr_DW_52.76.17.124;PG DB Size;CRITICAL;SOFT;2;NRPE: Unable to read output
[1452236705] SERVICE ALERT: credr_DW_52.76.17.124;PG Ideal Connections;CRITICAL;SOFT;2;POSTGRES_TXN_IDLE CRITICAL: DB "apps_data" Total idle in transaction: 72
[1452236743] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 140
[1452236757] SERVICE ALERT: credr_DW_52.76.17.124;PG DB Size;CRITICAL;HARD;3;NRPE: Unable to read output
[1452236765] SERVICE ALERT: credr_DW_52.76.17.124;PG Ideal Connections;OK;SOFT;3;POSTGRES_TXN_IDLE OK: DB "apps_data" not more than 50 idle in transaction
[1452236803] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 129
[1452236913] SERVICE ALERT: credr_DW_52.76.17.124;PG DB Size;OK;HARD;3;POSTGRES_DATABASE_SIZE OK: DB "apps_data" credr_dw: 762463304 (727 MB) apps_data: 480609352 (458 MB) apps_scrapping1: 406389832 (388 MB) apps_data_test: 402867272 (384 MB) test3: 289186888 (276 MB) test: 192062536 (183 MB) crm_replica: 95847496 (91 MB) solardb: 92996680 (89 MB) test2: 81675336 (78 MB) dashboard: 63964232 (61 MB) credr_analytics: 60425288 (58 MB) triggers_test: 43312200 (41 MB) credr_data: 40010824 (38 MB) crm_level1: 20841544 (20 MB) apps_internal: 8152136 (7961 kB) user_recommendation: 7218248 (7049 kB) crm_logging: 7144520 (6977 kB) credr_data_warehouse: 7029832 (6865 kB) postgres: 6923336 (6761 kB) redshift_testing: 6784072 (6625 kB) whatsapp: 6497352 (6345 kB) template1: 6448200 (6297 kB) template0: 6332932 (6185 kB)
[1452236982] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=12
[1452239082] SERVICE ALERT: Nagios Server;CPU Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_cpu_stats.sh, ...) failed. errno is 2: No such file or directory
[1452239135] SERVICE ALERT: Nagios Server;CPU Usage;UNKNOWN;SOFT;2;UNKNOWN: iostat not found or is not executable by the nagios user.
[1452239148] SERVICE ALERT: Nagios Server;CPU Usage;UNKNOWN;HARD;3;UNKNOWN: iostat not found or is not executable by the nagios user.
[1452239181] SERVICE ALERT: Nagios Server;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.50%
[1452240122] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 206 processes
[1452240182] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;SOFT;2;PROCS OK: 199 processes
[1452240362] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 213 processes
[1452240422] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 213 processes
[1452240482] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;HARD;3;PROCS WARNING: 206 processes
[1452241472] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.24% system=1.27%, iowait=0.00%, idle=56.49%, nice=0.00%, steal=0.00%
[1452241532] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=19.45% system=0.75%, iowait=0.00%, idle=79.80%, nice=0.00%, steal=0.00%
[1452242282] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;HARD;3;PROCS OK: 169 processes
[1452244152] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 111
[1452244213] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 111
[1452244272] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 111
[1452245531] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=89
[1452255123] HOST ALERT: credr_db_prod_slave_52.76.187.176;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1452255135] HOST ALERT: credr_db_prod_slave_52.76.187.176;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1452255138] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;CPU Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255189] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255198] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;CPU Usage;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255204] HOST ALERT: credr_db_prod_slave_52.76.187.176;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1452255206] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Ping Check;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255207] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255222] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255223] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255236] HOST ALERT: credr_db_prod_slave_52.76.187.176;UP;SOFT;3;PING WARNING - Packet loss = 96%, RTA = 0.55 ms
[1452255237] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Current Load;OK;SOFT;2;OK - load average: 0.00, 0.01, 0.05
[1452255237] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.50%, idle=99.00%, nice=0.00%, steal=0.00%
[1452255240] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Ping Check;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1452255249] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1452255252] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Root Partition;OK;SOFT;2;DISK OK - free space: / 92517 MB (95% inode=98%):
[1452255253] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Current Users;OK;SOFT;2;USERS OK - 1 users currently logged in
[1452255273] SERVICE ALERT: credr_db_prod_slave_52.76.187.176;Memory;OK;SOFT;3;[MEMORY] Total: 3951 MB - Used: 240 MB - 6% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1452264032] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=55.01% system=3.08%, iowait=0.00%, idle=41.90%, nice=0.00%, steal=0.00%
[1452264092] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=27.18% system=1.54%, iowait=0.00%, idle=71.28%, nice=0.00%, steal=0.00%
[1452264272] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.72% system=1.26%, iowait=0.00%, idle=55.03%, nice=0.00%, steal=0.00%
[1452264332] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=35.94% system=0.98%, iowait=0.24%, idle=62.84%, nice=0.00%, steal=0.00%