[1454869903] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=1.06% system=2.66%, iowait=63.30%, idle=31.91%, nice=0.00%, steal=1.06%
[1454869935] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=55.07% system=18.84%, iowait=24.64%, idle=0.97%, nice=0.00%, steal=0.48%
[1454869963] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=49.49% system=1.53%, iowait=0.00%, idle=48.47%, nice=0.00%, steal=0.51%
[1454869995] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=6.70% system=1.03%, iowait=2.06%, idle=89.69%, nice=0.00%, steal=0.52%
[1454874715] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=40.40% system=9.60%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454874775] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=43.78% system=6.47%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454874835] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454875195] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1454875915] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.06% system=6.44%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454875975] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.27% system=3.98%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.50%
[1454876035] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454876575] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454876575] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1454878912] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=4.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1454878972] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.27% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454879032] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=42.79% system=6.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1454879212] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454882632] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1454889587] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 230
[1454889648] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454889707] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454889887] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=84
[1454890067] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 213
[1454890127] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;SOFT;2;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1454893252] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.50% system=6.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454893312] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.78% system=5.47%, iowait=0.00%, idle=49.25%, nice=0.00%, steal=0.50%
[1454893372] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=44.50% system=5.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454893552] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.50%, idle=99.00%, nice=0.00%, steal=0.00%
[1454896972] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=5.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454897032] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.00% system=6.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454897092] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=47.00% system=3.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454897632] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.50%
[1454900512] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.72% system=6.53%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454900572] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.00% system=6.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454900632] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=41.00% system=9.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454900812] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=35.50% system=0.50%, iowait=0.00%, idle=63.50%, nice=0.00%, steal=0.50%
[1454904772] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454904832] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.27% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454904892] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454905432] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=39.70% system=10.05%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1454914453] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=57.07% system=0.51%, iowait=0.00%, idle=42.42%, nice=0.00%, steal=0.00%
[1454914513] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=9.90% system=1.49%, iowait=0.00%, idle=88.12%, nice=0.00%, steal=0.50%
[1454919003] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor credrWebApp;CRITICAL;SOFT;1;CRITICAL: Could not get status of credrWebApp
[1454919040] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor credrWebApp;CRITICAL;SOFT;2;CRITICAL: Could not get status of credrWebApp
[1454919100] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor credrWebApp;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1454923781] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1454923841] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1454923901] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1454924441] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;OK;HARD;3;USERS OK - 5 users currently logged in
[1454924716] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=4.02%, iowait=0.00%, idle=45.23%, nice=0.00%, steal=0.50%
[1454924776] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.11% system=3.11%, iowait=0.00%, idle=93.26%, nice=0.00%, steal=0.52%
[1454925161] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1454925221] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1454925281] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1454925821] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;OK;HARD;3;USERS OK - 5 users currently logged in
[1454925821] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;Current Users;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1454926287] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;Current Users;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1454926319] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credr-core;CRITICAL;SOFT;1;NRPE: Command 'check_celery_credrcore' not defined
[1454926379] SERVICE ALERT: credr_stage_52.76.16.57;Celery Process credr-core;OK;SOFT;2;credrcore worker process running: 7
[1454926688] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;Current Users;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1454926749] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;CRITICAL;SOFT;1;NRPE: Command 'check_celery_mem_credrcore' not defined
[1454926809] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;CRITICAL;SOFT;2;NRPE: Command 'check_celery_mem_credrcore' not defined
[1454926835] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory credr-core;OK;SOFT;3;Memory used by celery_credrcore process :: 141.75 MB
[1454927928] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.52% system=0.52%, iowait=47.94%, idle=51.03%, nice=0.00%, steal=0.00%
[1454927988] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454928881] SERVICE FLAPPING ALERT: credr_stage_52.76.16.57;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1454942319] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.50% system=1.00%, iowait=0.50%, idle=47.50%, nice=0.00%, steal=0.50%
[1454942379] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=46.97% system=0.00%, iowait=0.51%, idle=52.53%, nice=0.00%, steal=0.00%
[1454945978] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1454945993] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1454946038] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1454946053] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1454946098] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1454946113] SERVICE ALERT: credr_stage_52.76.16.57;Celery Memory Usage;OK;SOFT;3;Memory used by celery process :: 292.875 MB
[1454949336] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;OK;HARD;3;USERS OK - 4 users currently logged in