[1455129041] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.30% system=1.04%, iowait=0.00%, idle=47.15%, nice=0.00%, steal=0.52%
[1455129101] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=63.10% system=3.74%, iowait=0.53%, idle=32.62%, nice=0.00%, steal=0.00%
[1455129131] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.51% system=1.02%, iowait=0.00%, idle=47.96%, nice=0.00%, steal=0.51%
[1455129161] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=16.75% system=3.66%, iowait=59.16%, idle=19.90%, nice=0.00%, steal=0.52%
[1455129191] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.51% system=0.00%, iowait=0.00%, idle=48.99%, nice=0.00%, steal=0.51%
[1455129251] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.50%
[1455129341] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=4.50% system=1.00%, iowait=0.00%, idle=94.00%, nice=0.00%, steal=0.50%
[1455132221] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;SOFT;1;CRITICAL -  Error: Redis connection to localhost:6379 failed - connect ECONNREFUSED
[1455132280] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;OK;SOFT;2;OK - 511.13kB of memory used.
[1455134381] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.77% system=3.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1455134441] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=42.29% system=7.96%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455134501] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.27% system=4.98%, iowait=0.00%, idle=49.25%, nice=0.00%, steal=0.50%
[1455134681] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=38.19% system=1.51%, iowait=0.00%, idle=60.30%, nice=0.00%, steal=0.00%
[1455135221] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.00% system=7.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1455135281] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=46.27% system=3.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455135341] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=44.50% system=5.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455135521] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=39.70% system=10.05%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455135521] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1455138941] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1455148874] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1455148934] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;2;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 237
[1455148994] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 231
[1455149175] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=99
[1455165221] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455165281] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=4.04% system=2.02%, iowait=0.00%, idle=93.94%, nice=0.00%, steal=0.00%
[1455165641] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.00% system=4.50%, iowait=0.00%, idle=53.50%, nice=0.00%, steal=0.00%
[1455165701] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=38.69% system=5.03%, iowait=0.00%, idle=55.78%, nice=0.00%, steal=0.50%
[1455171821] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=6.00%, iowait=0.50%, idle=47.50%, nice=0.00%, steal=0.50%
[1455171881] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.77% system=4.48%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455171941] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=44.28% system=5.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1455172121] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.50%, iowait=0.00%, idle=99.00%, nice=0.00%, steal=0.00%
[1455172931] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.78% system=1.00%, iowait=1.49%, idle=53.73%, nice=0.00%, steal=0.00%
[1455172991] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.50%, idle=99.00%, nice=0.00%, steal=0.00%
[1455177941] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=69.68% system=2.66%, iowait=0.00%, idle=27.13%, nice=0.00%, steal=0.53%
[1455178001] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=67.71% system=1.04%, iowait=0.00%, idle=30.73%, nice=0.00%, steal=0.52%
[1455178061] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=68.53% system=2.54%, iowait=0.00%, idle=27.92%, nice=0.00%, steal=1.02%
[1455178421] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=29.74% system=0.00%, iowait=0.00%, idle=70.26%, nice=0.00%, steal=0.00%
[1455178601] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.26% system=1.51%, iowait=0.00%, idle=46.73%, nice=0.00%, steal=0.50%
[1455178661] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=55.39% system=1.96%, iowait=0.00%, idle=42.16%, nice=0.00%, steal=0.49%
[1455178721] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=33.50% system=0.00%, iowait=0.00%, idle=65.99%, nice=0.00%, steal=0.51%
[1455178901] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.23% system=2.26%, iowait=0.00%, idle=46.61%, nice=0.00%, steal=0.90%
[1455178961] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=44.62% system=3.23%, iowait=0.00%, idle=52.15%, nice=0.00%, steal=0.00%
[1455184361] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=67.17% system=2.02%, iowait=1.52%, idle=29.29%, nice=0.00%, steal=0.00%
[1455184421] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=52.76% system=2.51%, iowait=1.51%, idle=42.71%, nice=0.00%, steal=0.50%
[1455184481] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=19.58% system=3.70%, iowait=0.00%, idle=76.19%, nice=0.00%, steal=0.53%
[1455191900] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9998 MB (26% inode=97%):
[1455191960] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9997 MB (26% inode=97%):
[1455192020] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;3;DISK OK - free space: / 10009 MB (26% inode=97%):
[1455192380] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9992 MB (25% inode=97%):
[1455192440] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10004 MB (26% inode=97%):
[1455192620] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 10000 MB (26% inode=97%):
[1455192680] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9999 MB (26% inode=97%):
[1455192740] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;3;DISK OK - free space: / 10001 MB (26% inode=97%):
[1455192920] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9997 MB (26% inode=97%):
[1455192980] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9990 MB (25% inode=97%):
[1455193040] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;HARD;3;DISK WARNING - free space: / 9982 MB (25% inode=97%):
[1455195200] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;HARD;3;DISK OK - free space: / 16204 MB (42% inode=97%):