[1451110087] SERVICE ALERT: credr_stage_52.74.191.24;Nginx 80;CRITICAL;SOFT;1;connect to address 52.74.191.24 and port 80: Connection refused
[1451110107] SERVICE ALERT: credr_stage_52.74.191.24;Nginx 80;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 393 bytes in 0.001 second response time
[1451110745] SERVICE ALERT: credr_stage_52.74.191.24;Nginx 443;WARNING;SOFT;1;Generic error code.
[1451110820] SERVICE ALERT: credr_stage_52.74.191.24;Nginx 443;WARNING;SOFT;2;Generic error code.
[1451110826] SERVICE ALERT: credr_stage_52.74.191.24;Nginx 443;WARNING;HARD;3;Generic error code.
[1451111035] SERVICE ALERT: credr_stage_52.74.191.24;Nginx 443;OK;HARD;3;HTTP/1.1 200 OK
[1451116998] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1451117002] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1451117122] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;UNKNOWN;HARD;3;NRPE: Unable to read output
[1451117314] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1451118859] SERVICE ALERT: credr_stage_52.74.191.24;Supervisor;OK;HARD;3;credrWebApp                      RUNNING    pid 27979, uptime 0:34:17
[1451124700] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451124705] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451124825] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451124866] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451124985] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451124991] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451125101] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451127527] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.004 second response time
[1451127758] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;1;connect to address 52.74.244.220 and port 80: Connection refused
[1451127878] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;SOFT;2;connect to address 52.74.244.220 and port 80: Connection refused
[1451127998] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;CRITICAL;HARD;3;connect to address 52.74.244.220 and port 80: Connection refused
[1451127998] SERVICE FLAPPING ALERT: credr_dev_52.74.244.220;HTTP;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1451129198] SERVICE ALERT: credr_dev_52.74.244.220;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5800 bytes in 0.005 second response time
[1451133308] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1451133334] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;CRITICAL;SOFT;1;NRPE: Command 'check_cpu' not defined
[1451133348] SERVICE ALERT: credr_prod_52.74.247.15;Zombie Process;CRITICAL;SOFT;1;NRPE: Command 'check_zombie_procs' not defined
[1451133364] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1451133390] SERVICE ALERT: credr_prod_52.74.247.15;Current Load;CRITICAL;SOFT;1;NRPE: Command 'check_load' not defined
[1451133420] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1451133446] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;CRITICAL;SOFT;1;NRPE: Command 'check_users' not defined
[1451133454] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;CRITICAL;SOFT;2;NRPE: Command 'check_cpu' not defined
[1451133468] SERVICE ALERT: credr_prod_52.74.247.15;Zombie Process;CRITICAL;SOFT;2;NRPE: Command 'check_zombie_procs' not defined
[1451133476] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1451133502] SERVICE ALERT: credr_prod_52.74.247.15;Memory;CRITICAL;SOFT;1;NRPE: Command 'check_mem' not defined
[1451133510] SERVICE ALERT: credr_prod_52.74.247.15;Current Load;CRITICAL;SOFT;2;NRPE: Command 'check_load' not defined
[1451133532] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1451133566] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;CRITICAL;SOFT;2;NRPE: Command 'check_users' not defined
[1451133574] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;CRITICAL;HARD;3;NRPE: Command 'check_cpu' not defined
[1451133588] SERVICE ALERT: credr_prod_52.74.247.15;Zombie Process;CRITICAL;HARD;3;NRPE: Command 'check_zombie_procs' not defined
[1451133588] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1451133614] SERVICE ALERT: credr_prod_52.74.247.15;Ping Check;CRITICAL;SOFT;1;NRPE: Command 'check_ping' not defined
[1451133622] SERVICE ALERT: credr_prod_52.74.247.15;Memory;CRITICAL;SOFT;2;NRPE: Command 'check_mem' not defined
[1451133630] SERVICE ALERT: credr_prod_52.74.247.15;Current Load;CRITICAL;HARD;3;NRPE: Command 'check_load' not defined
[1451133644] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1451133670] SERVICE ALERT: credr_prod_52.74.247.15;Root Partition;CRITICAL;SOFT;1;NRPE: Command 'check_disk' not defined
[1451133686] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;CRITICAL;HARD;3;NRPE: Command 'check_users' not defined
[1451133700] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1451133738] SERVICE ALERT: credr_prod_52.74.247.15;Ping Check;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1451133742] SERVICE ALERT: credr_prod_52.74.247.15;Memory;UNKNOWN;HARD;3;NRPE: Unable to read output
[1451133756] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1451133782] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1451133790] SERVICE ALERT: credr_prod_52.74.247.15;Root Partition;OK;SOFT;2;DISK OK - free space: / 86340 MB (89% inode=95%):
[1451133812] HOST ALERT: credr_prod_52.74.247.15;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1451133902] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1451134022] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;UNKNOWN;HARD;3;NRPE: Unable to read output
[1451134174] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;UNKNOWN;HARD;3;NRPE: Unable to read output
[1451134188] SERVICE ALERT: credr_prod_52.74.247.15;Zombie Process;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1451134230] SERVICE ALERT: credr_prod_52.74.247.15;Current Load;OK;HARD;3;OK - load average: 1.52, 1.16, 1.02
[1451134286] SERVICE ALERT: credr_prod_52.74.247.15;Current Users;OK;HARD;3;USERS OK - 4 users currently logged in
[1451135542] SERVICE ALERT: credr_prod_52.74.247.15;Memory;OK;HARD;3;[MEMORY] Total: 16047 MB - Used: 13887 MB - 86% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1451135822] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor;CRITICAL;HARD;3;NRPE: Command 'check_super' not defined
[1451136082] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=46.19% system=1.05%, iowait=0.00%, idle=52.49%, nice=0.00%, steal=0.26%
[1451136682] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=26.80% system=1.24%, iowait=0.00%, idle=71.96%, nice=0.00%, steal=0.00%
[1451136682] SERVICE FLAPPING ALERT: credr_prod_52.74.247.15;CPU Usage;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1451138107] SERVICE FLAPPING ALERT: credr_prod_52.74.247.15;CPU Usage;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1451148079] SERVICE FLAPPING ALERT: credr_prod_52.74.247.15;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)