[1454439494] SERVICE ALERT: credr_prod_52.74.247.15;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1454439494] SERVICE ALERT: credr_prod_52.74.247.15;Nginx http;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1454439524] HOST ALERT: credr_prod_52.74.247.15;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454439544] SERVICE ALERT: credr_prod_52.74.247.15;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1454439544] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor Celery;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1454439554] SERVICE ALERT: credr_prod_52.74.247.15;Nginx http;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1454439554] SERVICE ALERT: credr_prod_52.74.247.15;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1454439571] HOST ALERT: credr_prod_52.74.247.15;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.50 ms
[1454439574] SERVICE ALERT: credr_prod_52.74.247.15;Zombie Process;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1454439574] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1454439574] SERVICE ALERT: credr_prod_52.74.247.15;Zombie Process;OK;SOFT;2;PROCS OK: 1 process with STATE = Z
[1454439574] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;2;ERROR: CANNOT RETRIEVE URL: 503 Service Unavailable: Back-end server is at capacity
[1454439574] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor WebApp;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1454439574] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor WebApp;OK;SOFT;2;credrWebApp                      RUNNING    pid 1590, uptime 0:00:01
[1454439575] SERVICE ALERT: credr_prod_52.74.247.15;Memory;OK;SOFT;2;[MEMORY] Total: 32175 MB - Used: 609 MB - 1% [SWAP] Total: 0 MB - Used: 0 MB - 0%
[1454439575] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor Celery;CRITICAL;SOFT;2;celery_smsserv                   BACKOFF    Exited too quickly (process log may have details)
[1454439604] SERVICE ALERT: credr_prod_52.74.247.15;Nginx http;OK;SOFT;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 374 bytes in 0.001 second response time
[1454439604] SERVICE ALERT: credr_prod_52.74.247.15;SSH;OK;SOFT;3;SSH OK - OpenSSH_6.6.1p1 Ubuntu-2ubuntu2.4 (protocol 2.0)
[1454439604] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor Celery;CRITICAL;HARD;3;celery_smsserv                   FATAL      Exited too quickly (process log may have details)
[1454439634] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;HARD;3;ERROR: CANNOT RETRIEVE URL: 503 Service Unavailable: Back-end server is at capacity
[1454439814] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;HARD;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1454439934] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;1;PROCS WARNING: 221 processes
[1454439964] SERVICE ALERT: credr_prod_52.74.247.15;Supervisor Celery;OK;HARD;3;celery_smsserv                   RUNNING    pid 2258, uptime 0:02:27
[1454439994] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;SOFT;2;PROCS WARNING: 224 processes
[1454440054] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;WARNING;HARD;3;PROCS WARNING: 219 processes
[1454440774] SERVICE ALERT: credr_prod_52.74.247.15;Total Processes;OK;HARD;3;PROCS OK: 213 processes
[1454457665] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 103
[1454457724] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;SOFT;2;POSTGRES_LOCKS OK: DB "apps_data" total=94
[1454480586] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=20.79% system=30.69%, iowait=0.50%, idle=47.52%, nice=0.00%, steal=0.50%
[1454480646] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=19.50% system=30.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454480706] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=21.39% system=28.36%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1454480886] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=57.71% system=28.36%, iowait=1.00%, idle=12.44%, nice=0.00%, steal=0.50%
[1454480946] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454481666] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=23.38% system=31.84%, iowait=0.00%, idle=44.78%, nice=0.00%, steal=0.00%
[1454481726] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=21.00% system=29.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454482626] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=20.00% system=30.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454482686] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=22.11% system=27.64%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1454485746] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=17.91% system=32.34%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454485806] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=21.11% system=32.66%, iowait=0.00%, idle=46.23%, nice=0.00%, steal=0.00%
[1454485866] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=20.90% system=28.86%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1454486586] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=39.00% system=59.50%, iowait=0.50%, idle=1.00%, nice=0.00%, steal=0.00%
[1454486646] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=38.12% system=61.39%, iowait=0.00%, idle=0.00%, nice=0.00%, steal=0.50%
[1454486706] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=34.65% system=64.36%, iowait=0.00%, idle=0.50%, nice=0.00%, steal=0.50%
[1454488686] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=27.50% system=27.00%, iowait=0.00%, idle=45.50%, nice=0.00%, steal=0.00%
[1454494626] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=35.50% system=33.00%, iowait=0.50%, idle=30.50%, nice=0.00%, steal=0.50%
[1454494686] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=22.61% system=27.14%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1454499006] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=66.00% system=31.00%, iowait=0.00%, idle=2.50%, nice=0.00%, steal=0.50%
[1454499066] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=27.23% system=22.77%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454499246] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=19.60% system=30.15%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1454499306] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=20.00% system=30.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454499366] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=20.79% system=31.68%, iowait=0.00%, idle=47.03%, nice=0.00%, steal=0.50%
[1454499396] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=40.00% system=1.00%, iowait=2.50%, idle=56.50%, nice=0.00%, steal=0.00%
[1454499456] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=2.54% system=0.00%, iowait=1.02%, idle=96.45%, nice=0.00%, steal=0.00%
[1454499636] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.50% system=2.50%, iowait=2.50%, idle=50.50%, nice=0.00%, steal=0.00%
[1454499696] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.51% system=0.00%, iowait=1.01%, idle=97.49%, nice=0.00%, steal=0.00%
[1454500266] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=21.78% system=28.22%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454500446] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=19.90% system=30.85%, iowait=0.00%, idle=48.76%, nice=0.00%, steal=0.50%
[1454500506] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=20.90% system=34.33%, iowait=0.50%, idle=43.78%, nice=0.00%, steal=0.50%
[1454500566] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=21.50% system=30.00%, iowait=0.00%, idle=48.50%, nice=0.00%, steal=0.00%
[1454501826] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=21.00% system=29.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454501826] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1454502006] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=33.33% system=34.33%, iowait=1.00%, idle=30.85%, nice=0.00%, steal=0.50%
[1454502066] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=31.68% system=34.65%, iowait=0.00%, idle=33.17%, nice=0.00%, steal=0.50%
[1454502126] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=29.85% system=36.82%, iowait=0.50%, idle=32.84%, nice=0.00%, steal=0.00%
[1454502666] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=19.31% system=30.69%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454505726] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)