[1454697186] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.51% system=0.00%, iowait=0.51%, idle=48.48%, nice=0.00%, steal=0.51%
[1454697246] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454700363] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.22% system=4.52%, iowait=0.00%, idle=52.26%, nice=0.00%, steal=0.00%
[1454700423] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454700483] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=21.00% system=5.00%, iowait=0.00%, idle=74.00%, nice=0.00%, steal=0.00%
[1454702823] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454702883] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=21.50% system=1.50%, iowait=0.00%, idle=77.00%, nice=0.00%, steal=0.00%
[1454703063] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.05% system=4.95%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454703123] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=34.00% system=4.50%, iowait=0.00%, idle=61.00%, nice=0.00%, steal=0.50%
[1454704383] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454704443] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454716864] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454716924] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 100
[1454716984] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 217
[1454717524] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1454720164] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9999 MB (26% inode=97%):
[1454720224] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10016 MB (26% inode=97%):
[1454722744] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9997 MB (26% inode=97%):
[1454722804] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10011 MB (26% inode=97%):
[1454724064] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9998 MB (26% inode=97%):
[1454724124] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10007 MB (26% inode=97%):
[1454724664] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9998 MB (26% inode=97%):
[1454724724] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10006 MB (26% inode=97%):
[1454725264] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9997 MB (26% inode=97%):
[1454725324] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10003 MB (26% inode=97%):
[1454725864] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9995 MB (26% inode=97%):
[1454725924] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;SOFT;2;DISK OK - free space: / 10001 MB (26% inode=97%):
[1454726104] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9985 MB (25% inode=97%):
[1454726164] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9991 MB (25% inode=97%):
[1454726224] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;WARNING;HARD;3;DISK WARNING - free space: / 10000 MB (26% inode=97%):
[1454736303] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=0.50%, iowait=0.50%, idle=49.00%, nice=0.00%, steal=0.00%
[1454736363] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.58%, iowait=0.00%, idle=98.84%, nice=0.00%, steal=0.58%
[1454736723] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.50% system=0.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454736783] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.00% system=0.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454736843] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=50.00% system=0.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454737023] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=1.00%, iowait=0.00%, idle=98.50%, nice=0.00%, steal=0.00%
[1454739003] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=49.75% system=0.00%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1454739063] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454739364] SERVICE ALERT: credr_prodsolr_52.74.114.132;Root Partition;OK;HARD;3;DISK OK - free space: / 13370 MB (34% inode=97%):
[1454739783] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.75% system=1.00%, iowait=0.00%, idle=47.76%, nice=0.00%, steal=0.50%
[1454739843] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454740383] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454740443] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454741166] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=1.01%, iowait=0.50%, idle=48.24%, nice=0.00%, steal=0.00%
[1454741226] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=46.97% system=0.00%, iowait=0.51%, idle=52.53%, nice=0.00%, steal=0.00%
[1454741343] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=47.74% system=5.03%, iowait=0.00%, idle=47.24%, nice=0.00%, steal=0.00%
[1454741403] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.73% system=3.52%, iowait=0.00%, idle=50.75%, nice=0.00%, steal=0.00%
[1454741463] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454741643] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.76% system=0.00%, iowait=0.00%, idle=48.24%, nice=0.00%, steal=0.00%
[1454741703] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=99.42%, nice=0.00%, steal=0.58%
[1454741883] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454741943] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454742003] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454742183] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1454742243] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454742783] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454742843] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1454743203] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.00% system=1.50%, iowait=0.00%, idle=47.50%, nice=0.00%, steal=0.00%
[1454743263] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.00% system=0.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454743323] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1454744093] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor WebApp;CRITICAL;SOFT;1;CRITICAL: Could not get status of credrWebApp
[1454744125] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1454744153] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor WebApp;CRITICAL;SOFT;2;CRITICAL: Could not get status of credrWebApp
[1454744185] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1454744213] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor WebApp;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1454744222] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celery;CRITICAL;SOFT;1;CRITICAL: Could not get status of celery_smsserv
[1454744245] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;UNKNOWN;HARD;3;NRPE: Unable to read output
[1454744282] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celery;CRITICAL;SOFT;2;CRITICAL: Could not get status of celery_smsserv
[1454744342] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celery;CRITICAL;HARD;3;CRITICAL: Could not get status of celery_smsserv
[1454744514] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=3.59% system=3.59%, iowait=0.00%, idle=91.79%, nice=0.00%, steal=1.03%
[1454745936] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=57.44% system=1.03%, iowait=0.00%, idle=41.54%, nice=0.00%, steal=0.00%
[1454745996] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=28.23% system=2.39%, iowait=0.00%, idle=68.90%, nice=0.00%, steal=0.48%
[1454750856] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=52.79% system=1.52%, iowait=0.00%, idle=45.18%, nice=0.00%, steal=0.51%
[1454750916] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=42.93% system=2.09%, iowait=0.00%, idle=54.45%, nice=0.00%, steal=0.52%
[1454754696] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=90.16% system=1.04%, iowait=0.00%, idle=8.81%, nice=0.00%, steal=0.00%
[1454754807] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=32.47% system=0.52%, iowait=0.00%, idle=65.98%, nice=0.00%, steal=1.03%
[1454756246] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celery Credrcore;CRITICAL;SOFT;1;CRITICAL: Could not get status of celery_credrcore
[1454756250] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celery Credrcore;CRITICAL;SOFT;2;CRITICAL: Could not get status of celery_credrcore
[1454756310] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celery Credrcore;CRITICAL;HARD;3;CRITICAL: Could not get status of celery_credrcore
[1454756324] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor credrWebApp;CRITICAL;SOFT;1;CRITICAL: Could not get status of credrWebApp
[1454756328] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor credrWebApp;CRITICAL;SOFT;2;CRITICAL: Could not get status of credrWebApp
[1454756352] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celerybeat Credrcore;CRITICAL;SOFT;1;CRITICAL: Could not get status of celerybeat_credrcore
[1454756388] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor credrWebApp;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1454756412] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celerybeat Credrcore;CRITICAL;SOFT;2;CRITICAL: Could not get status of celerybeat_credrcore
[1454756472] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celerybeat Credrcore;CRITICAL;HARD;3;CRITICAL: Could not get status of celerybeat_credrcore
[1454761908] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celerybeat Credrcore;CRITICAL;SOFT;1;CRITICAL: Could not get status of celerybeat_credrcore
[1454761912] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celerybeat Credrcore;CRITICAL;SOFT;2;CRITICAL: Could not get status of celerybeat_credrcore
[1454761916] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor Celerybeat Credrcore;CRITICAL;HARD;3;CRITICAL: Could not get status of celerybeat_credrcore
[1454763006] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.00% system=1.02%, iowait=1.02%, idle=47.96%, nice=0.00%, steal=0.00%
[1454763066] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=47.47% system=0.00%, iowait=0.51%, idle=52.02%, nice=0.00%, steal=0.00%
[1454765643] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=49.18% system=4.92%, iowait=2.19%, idle=42.62%, nice=0.00%, steal=1.09%
[1454765703] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=15.23% system=5.08%, iowait=2.03%, idle=77.66%, nice=0.00%, steal=0.00%
[1454765762] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454765822] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;SOFT;2;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1454767566] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=0.50%, iowait=0.50%, idle=48.74%, nice=0.00%, steal=0.00%
[1454767626] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=47.00% system=0.00%, iowait=0.50%, idle=52.50%, nice=0.00%, steal=0.00%