[1452106712] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 300 processes
[1452106772] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;2;PROCS OK: 150 processes
[1452107672] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 328 processes
[1452107732] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 373 processes
[1452107792] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;HARD;3;PROCS WARNING: 235 processes
[1452112472] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 362 processes
[1452114452] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;HARD;3;PROCS WARNING: 297 processes
[1452134764] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=81.36% system=10.73%, iowait=0.00%, idle=7.91%, nice=0.00%, steal=0.00%
[1452134824] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=78.71% system=12.90%, iowait=0.00%, idle=7.74%, nice=0.00%, steal=0.65%
[1452134884] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=74.23% system=15.46%, iowait=0.00%, idle=9.79%, nice=0.00%, steal=0.52%
[1452135064] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=29.00% system=7.00%, iowait=0.00%, idle=64.00%, nice=0.00%, steal=0.00%
[1452138304] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.00% system=8.00%, iowait=0.00%, idle=41.00%, nice=0.00%, steal=0.00%
[1452138364] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=44.78% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1452145474] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.90% system=1.03%, iowait=0.00%, idle=48.07%, nice=0.00%, steal=0.00%
[1452145534] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=16.20% system=0.25%, iowait=0.00%, idle=83.54%, nice=0.00%, steal=0.00%
[1452145952] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;HARD;3;PROCS OK: 216 processes
[1452158134] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.03% system=0.76%, iowait=0.00%, idle=57.22%, nice=0.00%, steal=0.00%
[1452158194] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=16.36% system=0.52%, iowait=0.00%, idle=82.60%, nice=0.00%, steal=0.52%
[1452158374] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.50% system=2.25%, iowait=0.00%, idle=55.00%, nice=0.00%, steal=0.25%
[1452158434] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=21.43% system=0.77%, iowait=0.00%, idle=77.81%, nice=0.00%, steal=0.00%
[1452159334] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=52.71% system=1.55%, iowait=0.00%, idle=45.74%, nice=0.00%, steal=0.00%
[1452159394] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=18.32% system=0.25%, iowait=0.00%, idle=81.42%, nice=0.00%, steal=0.00%
[1452163052] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 286 processes
[1452163112] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;2;PROCS OK: 239 processes
[1452163832] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 329 processes
[1452163892] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;2;PROCS WARNING: 325 processes
[1452163952] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;HARD;3;PROCS WARNING: 318 processes
[1452164132] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;HARD;3;PROCS OK: 234 processes
[1452172052] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 266 processes
[1452172112] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;2;PROCS OK: 229 processes
[1452177512] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 317 processes
[1452177572] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;2;PROCS WARNING: 276 processes
[1452177632] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;3;PROCS OK: 222 processes
[1452177992] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;WARNING;SOFT;1;PROCS WARNING: 261 processes
[1452178052] SERVICE ALERT: credr_DW_52.76.17.124;Total Processes;OK;SOFT;2;PROCS OK: 222 processes
[1452182794] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.60% system=1.79%, iowait=0.00%, idle=55.61%, nice=0.00%, steal=0.00%
[1452182854] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=11.53% system=0.50%, iowait=0.00%, idle=87.97%, nice=0.00%, steal=0.00%