[1456338629] SERVICE ALERT: credr_prod_52.74.247.15;Nginx http;CRITICAL;SOFT;2;connect to address 52.74.247.15 and port 80: Connection refused
[1456338659] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;HARD;3;ERROR: CANNOT RETRIEVE URL: 503 Service Unavailable: Back-end server is at capacity
[1456338689] SERVICE ALERT: credr_prod_52.74.247.15;Nginx http;OK;SOFT;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 374 bytes in 0.001 second response time
[1456338751] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=60.22% system=2.21%, iowait=0.55%, idle=37.02%, nice=0.00%, steal=0.00%
[1456338779] SERVICE ALERT: credr_stage_52.76.16.57;Website Check;OK;HARD;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1456338811] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=7.33% system=0.52%, iowait=0.00%, idle=91.62%, nice=0.00%, steal=0.52%
[1456338839] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;HARD;3;CONTENT OK: WEBSITE IS RUNNING FINE
[1456358309] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 262
[1456358370] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 115
[1456358429] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 136
[1456359149] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=28
[1456362154] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=7.53% system=1.61%, iowait=38.71%, idle=52.15%, nice=0.00%, steal=0.00%
[1456362215] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.00% system=0.52%, iowait=51.55%, idle=47.42%, nice=0.00%, steal=0.52%
[1456362274] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=7.81% system=1.56%, iowait=38.02%, idle=52.08%, nice=0.00%, steal=0.52%
[1456362454] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=0.00% system=0.51%, iowait=48.47%, idle=51.02%, nice=0.00%, steal=0.00%
[1456369229] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1456369290] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 108
[1456369349] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 109
[1456369653] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=3.11% system=1.04%, iowait=39.90%, idle=55.44%, nice=0.00%, steal=0.52%
[1456369834] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=0.00% system=0.00%, iowait=52.02%, idle=47.47%, nice=0.00%, steal=0.51%
[1456372532] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=0.00% system=0.50%, iowait=39.00%, idle=60.50%, nice=0.00%, steal=0.00%
[1456372713] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=1.01% system=1.01%, iowait=44.72%, idle=52.26%, nice=0.00%, steal=1.01%
[1456372713] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1456376131] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1456380510] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1456380631] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=5.47% system=1.49%, iowait=2.99%, idle=90.05%, nice=0.00%, steal=0.00%
[1456380991] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=0.00%, iowait=43.43%, idle=56.57%, nice=0.00%, steal=0.00%
[1456381051] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.51% system=0.00%, iowait=47.69%, idle=51.79%, nice=0.00%, steal=0.00%
[1456381111] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456381292] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=0.50%, iowait=49.25%, idle=50.25%, nice=0.00%, steal=0.00%
[1456381351] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.00% system=0.50%, iowait=48.74%, idle=50.25%, nice=0.00%, steal=0.50%
[1456381411] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=0.00% system=0.50%, iowait=49.50%, idle=50.00%, nice=0.00%, steal=0.00%
[1456382312] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456382491] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=9.05% system=4.02%, iowait=77.89%, idle=8.54%, nice=0.00%, steal=0.50%
[1456382551] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=0.00% system=0.51%, iowait=48.99%, idle=50.51%, nice=0.00%, steal=0.00%
[1456382611] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=0.00% system=0.00%, iowait=49.75%, idle=49.75%, nice=0.00%, steal=0.50%
[1456382611] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1456384051] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456384412] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=0.00% system=0.00%, iowait=42.71%, idle=57.29%, nice=0.00%, steal=0.00%
[1456384471] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=3.55% system=0.51%, iowait=38.58%, idle=57.36%, nice=0.00%, steal=0.00%
[1456384532] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=0.00% system=0.00%, iowait=46.00%, idle=53.50%, nice=0.00%, steal=0.50%
[1456384891] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=14.50% system=1.50%, iowait=0.00%, idle=84.00%, nice=0.00%, steal=0.00%
[1456388131] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1456407121] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=40.40% system=1.52%, iowait=2.02%, idle=56.06%, nice=0.00%, steal=0.00%
[1456407181] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=2.49% system=0.00%, iowait=1.49%, idle=96.02%, nice=0.00%, steal=0.00%