[1452729686] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 103
[1452729745] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;2;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 217
[1452729805] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;HARD;3;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 226
[1452729985] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 185
[1452730165] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1452739583] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=86.63% system=10.47%, iowait=0.00%, idle=2.33%, nice=0.00%, steal=0.58%
[1452739642] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=85.81% system=10.97%, iowait=0.00%, idle=3.23%, nice=0.00%, steal=0.00%
[1452739702] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=66.67% system=7.27%, iowait=0.00%, idle=25.45%, nice=0.00%, steal=0.61%
[1452739882] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.50%, idle=99.00%, nice=0.00%, steal=0.00%
[1452762709] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=60.62% system=4.15%, iowait=0.00%, idle=35.23%, nice=0.00%, steal=0.00%
[1452762769] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=6.50% system=0.50%, iowait=0.00%, idle=92.50%, nice=0.00%, steal=0.50%
[1452764662] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.88% system=1.31%, iowait=0.26%, idle=56.54%, nice=0.00%, steal=0.00%
[1452764722] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=23.35% system=1.27%, iowait=0.25%, idle=75.13%, nice=0.00%, steal=0.00%
[1452767422] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.62% system=2.01%, iowait=0.00%, idle=51.38%, nice=0.00%, steal=0.00%
[1452767482] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=19.23% system=1.03%, iowait=0.00%, idle=79.74%, nice=0.00%, steal=0.00%
[1452771982] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.18% system=2.41%, iowait=0.00%, idle=56.15%, nice=0.00%, steal=0.27%
[1452772042] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=18.69% system=0.51%, iowait=0.25%, idle=80.56%, nice=0.00%, steal=0.00%
[1452777262] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=53.35% system=1.80%, iowait=0.26%, idle=44.59%, nice=0.00%, steal=0.00%
[1452777322] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=12.76% system=0.26%, iowait=0.00%, idle=86.98%, nice=0.00%, steal=0.00%
[1452777502] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=73.71% system=2.58%, iowait=0.00%, idle=23.71%, nice=0.00%, steal=0.00%
[1452777562] SERVICE ALERT: credr_prod_52.74.247.15;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=21.54% system=1.28%, iowait=0.00%, idle=77.18%, nice=0.00%, steal=0.00%