[1456185630] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 134
[1456185689] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 129
[1456185749] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 127
[1456186290] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=28
[1456189235] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.54% system=4.95%, iowait=0.99%, idle=48.51%, nice=0.00%, steal=0.00%
[1456189295] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.78% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1456189355] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.78% system=5.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1456189715] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456192955] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.50% system=4.00%, iowait=0.50%, idle=49.50%, nice=0.00%, steal=0.50%
[1456193015] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=46.50% system=3.50%, iowait=0.50%, idle=49.50%, nice=0.00%, steal=0.00%
[1456193075] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.28% system=6.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1456193795] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456196495] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.50% system=8.00%, iowait=0.00%, idle=50.50%, nice=0.00%, steal=0.00%
[1456196555] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.00% system=4.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1456196615] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.28% system=6.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1456197515] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456209215] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=56.85% system=3.05%, iowait=12.18%, idle=27.41%, nice=0.00%, steal=0.51%
[1456209275] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456211975] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=41.71% system=8.04%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1456212035] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.50% system=5.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456212095] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=41.29% system=8.96%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1456212275] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=39.00% system=11.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456212455] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=44.22% system=6.03%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1456212515] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=41.00% system=9.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456212575] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.07% system=6.93%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.50%
[1456212755] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456212755] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1456213655] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.08% system=7.43%, iowait=0.50%, idle=49.50%, nice=0.00%, steal=0.50%
[1456213715] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.50% system=5.50%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1456213775] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=54.50% system=12.00%, iowait=0.50%, idle=32.50%, nice=0.00%, steal=0.50%
[1456214495] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1456217375] SERVICE FLAPPING ALERT: credr_DW_52.76.17.124;CPU Usage;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1456221335] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=67.38% system=30.48%, iowait=0.00%, idle=1.60%, nice=0.00%, steal=0.53%
[1456221395] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=72.83% system=25.00%, iowait=0.00%, idle=2.17%, nice=0.00%, steal=0.00%
[1456221455] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=69.68% system=28.19%, iowait=0.00%, idle=1.60%, nice=0.00%, steal=0.53%
[1456221815] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=28.36% system=24.88%, iowait=0.00%, idle=46.27%, nice=0.00%, steal=0.50%
[1456223911] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.78% system=0.00%, iowait=0.51%, idle=47.72%, nice=0.00%, steal=0.00%
[1456223971] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=7.00% system=2.00%, iowait=0.00%, idle=91.00%, nice=0.00%, steal=0.00%
[1456230661] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=59.30% system=3.27%, iowait=0.00%, idle=37.44%, nice=0.00%, steal=0.00%
[1456230721] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=74.37% system=3.55%, iowait=0.00%, idle=22.08%, nice=0.00%, steal=0.00%
[1456230781] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=48.99% system=3.28%, iowait=0.00%, idle=47.47%, nice=0.00%, steal=0.25%
[1456231141] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=74.06% system=3.27%, iowait=0.00%, idle=22.67%, nice=0.00%, steal=0.00%
[1456232039] SERVICE ALERT: credr_dev_52.74.244.220;Zombie Process;WARNING;SOFT;1;PROCS WARNING: 6 processes with STATE = Z
[1456232099] SERVICE ALERT: credr_dev_52.74.244.220;Zombie Process;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1456233301] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=1.25% system=0.25%, iowait=0.00%, idle=98.50%, nice=0.00%, steal=0.00%
[1456235101] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=79.55% system=1.00%, iowait=0.00%, idle=19.20%, nice=0.00%, steal=0.25%
[1456235161] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.25%, iowait=0.00%, idle=99.75%, nice=0.00%, steal=0.00%