[1454956273] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=59.59% system=2.07%, iowait=0.52%, idle=37.31%, nice=0.00%, steal=0.52%
[1454956333] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=29.41% system=8.02%, iowait=25.13%, idle=36.90%, nice=0.00%, steal=0.53%
[1454957106] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor WebApp;CRITICAL;SOFT;1;CRITICAL: Could not get status of credrWebApp
[1454957165] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor WebApp;CRITICAL;SOFT;2;CRITICAL: Could not get status of credrWebApp
[1454957225] SERVICE ALERT: credr_stage_52.76.16.57;Supervisor WebApp;CRITICAL;HARD;3;CRITICAL: Could not get status of credrWebApp
[1454960164] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;CRITICAL;SOFT;1;CRITICAL -  Error: Redis connection to localhost:6379 failed - connect ECONNREFUSED
[1454960224] SERVICE ALERT: credr_stage_52.76.16.57;Redis Memory Usage;OK;SOFT;2;OK - 511.13kB of memory used.
[1454975976] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 230
[1454976036] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454976096] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 106
[1454976277] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=84
[1454976456] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;CRITICAL;SOFT;1;POSTGRES_LOCKS CRITICAL: DB "apps_data" total locks: 213
[1454976516] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;SOFT;2;POSTGRES_LOCKS OK: DB "apps_data" total=1
[1454989904] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=47.00% system=3.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454989964] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.52% system=0.00%, iowait=0.00%, idle=96.48%, nice=0.00%, steal=0.00%
[1454991044] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454991104] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=43.78% system=5.97%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1454991164] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=28.50% system=4.50%, iowait=0.00%, idle=67.00%, nice=0.00%, steal=0.00%
[1454991344] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.27% system=4.48%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1454991404] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=36.00% system=3.50%, iowait=0.00%, idle=60.50%, nice=0.00%, steal=0.00%
[1454991584] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.00% system=4.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1454991644] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.48% system=0.50%, iowait=0.00%, idle=95.52%, nice=0.00%, steal=0.50%
[1454992004] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=46.27% system=3.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454992064] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=45.27% system=4.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.00%
[1454992124] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.00% system=5.50%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1454992304] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=3.52% system=0.50%, iowait=0.00%, idle=95.98%, nice=0.00%, steal=0.00%
[1454998898] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=79.59% system=1.02%, iowait=0.00%, idle=18.88%, nice=0.00%, steal=0.51%
[1454998958] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=82.83% system=0.00%, iowait=0.51%, idle=16.16%, nice=0.00%, steal=0.51%
[1454999018] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=84.42% system=0.50%, iowait=0.00%, idle=15.08%, nice=0.00%, steal=0.00%
[1454999378] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=6.77% system=0.00%, iowait=1.04%, idle=91.67%, nice=0.00%, steal=0.52%
[1454999739] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1454999799] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1454999859] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1455000759] SERVICE ALERT: credr_stage_52.76.16.57;Current Users;OK;HARD;3;USERS OK - 6 users currently logged in
[1455003698] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.75% system=1.01%, iowait=1.01%, idle=47.24%, nice=0.00%, steal=0.00%
[1455003758] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=45.96% system=0.00%, iowait=1.01%, idle=52.53%, nice=0.00%, steal=0.51%
[1455004497] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=49.21% system=10.58%, iowait=0.00%, idle=39.68%, nice=0.00%, steal=0.53%
[1455004557] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=2.01% system=3.52%, iowait=0.00%, idle=93.97%, nice=0.00%, steal=0.50%
[1455004917] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=55.87% system=7.26%, iowait=0.00%, idle=35.75%, nice=0.00%, steal=1.12%
[1455004977] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=3.03% system=2.02%, iowait=0.00%, idle=94.44%, nice=0.00%, steal=0.51%
[1455007358] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=66.67% system=4.04%, iowait=0.51%, idle=28.28%, nice=0.00%, steal=0.51%
[1455007418] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=66.33% system=3.02%, iowait=0.50%, idle=30.15%, nice=0.00%, steal=0.00%
[1455007438] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1455007440] SERVICE ALERT: credr_stage_52.76.16.57;Website Check;UNKNOWN;SOFT;1;ERROR: CANNOT RETRIEVE URL: 502 Bad Gateway
[1455007478] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=61.81% system=3.52%, iowait=0.50%, idle=34.17%, nice=0.00%, steal=0.00%
[1455007498] SERVICE ALERT: credr_prod_52.74.247.15;Website Check;OK;SOFT;2;CONTENT OK: WEBSITE IS RUNNING FINE
[1455007500] SERVICE ALERT: credr_stage_52.76.16.57;Website Check;OK;SOFT;2;CONTENT OK: WEBSITE IS RUNNING FINE
[1455007658] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=45.73% system=0.50%, iowait=0.50%, idle=53.27%, nice=0.00%, steal=0.00%
[1455009818] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.25% system=1.01%, iowait=0.50%, idle=48.24%, nice=0.00%, steal=0.00%
[1455009878] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=46.31% system=0.99%, iowait=0.99%, idle=51.23%, nice=0.00%, steal=0.49%
[1455010958] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=51.26% system=0.50%, iowait=1.01%, idle=47.24%, nice=0.00%, steal=0.00%
[1455011018] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=47.21% system=0.00%, iowait=0.51%, idle=52.28%, nice=0.00%, steal=0.00%
[1455029637] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.30% system=7.73%, iowait=0.00%, idle=48.45%, nice=0.00%, steal=0.52%
[1455029697] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=7.25% system=7.25%, iowait=0.00%, idle=84.97%, nice=0.00%, steal=0.52%