[1450851616] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 2 users currently logged in
[1450851676] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 2 users currently logged in
[1450851736] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 3 users currently logged in
[1450851796] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 3 users currently logged in
[1450852857] SERVICE ALERT: Nagios Server;Current Users;WARNING;SOFT;1;USERS WARNING - 3 users currently logged in
[1450852917] SERVICE ALERT: Nagios Server;Current Users;WARNING;SOFT;2;USERS WARNING - 3 users currently logged in
[1450852977] SERVICE ALERT: Nagios Server;Current Users;WARNING;SOFT;3;USERS WARNING - 3 users currently logged in
[1450853037] SERVICE ALERT: Nagios Server;Current Users;WARNING;HARD;4;USERS WARNING - 3 users currently logged in
[1450853072] SERVICE ALERT: Nagios Server;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 94 processes with STATE = RSZDT
[1450853132] SERVICE ALERT: Nagios Server;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 94 processes with STATE = RSZDT
[1450853192] SERVICE ALERT: Nagios Server;Total Processes;CRITICAL;SOFT;3;PROCS CRITICAL: 90 processes with STATE = RSZDT
[1450853252] SERVICE ALERT: Nagios Server;Total Processes;CRITICAL;HARD;4;PROCS CRITICAL: 89 processes with STATE = RSZDT
[1450855436] SERVICE ALERT: Nagios Server;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1450866835] SERVICE ALERT: Nagios Server;Current Users;WARNING;SOFT;1;USERS WARNING - 2 users currently logged in
[1450866895] SERVICE ALERT: Nagios Server;Current Users;WARNING;SOFT;2;USERS WARNING - 2 users currently logged in
[1450866955] SERVICE ALERT: Nagios Server;Current Users;WARNING;SOFT;3;USERS WARNING - 2 users currently logged in
[1450867015] SERVICE ALERT: Nagios Server;Current Users;WARNING;HARD;4;USERS WARNING - 2 users currently logged in
[1450868896] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;CRITICAL;SOFT;1;NRPE: Command 'check_ping' not defined
[1450868963] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 1177 MB (15% inode=62%):
[1450869016] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;CRITICAL;SOFT;2;NRPE: Command 'check_ping' not defined
[1450869082] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 1177 MB (15% inode=62%):
[1450869134] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;CRITICAL;HARD;3;NRPE: Command 'check_ping' not defined
[1450869201] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;WARNING;HARD;3;DISK WARNING - free space: / 1177 MB (15% inode=62%):
[1450869527] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;UNKNOWN;HARD;3;You must specify a server address or host name
[1450870083] SERVICE ALERT: credr_dev_52.74.244.220;Ping Check;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1450870203] SERVICE ALERT: credr_dev_52.74.244.220;Root Partition;OK;HARD;3;DISK OK - free space: / 1177 MB (15% inode=62%):
[1450872070] SERVICE ALERT: credr_dev_52.74.244.220;SSH;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1450872076] SERVICE ALERT: credr_dev_52.74.244.220;SSH;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1450872087] SERVICE ALERT: credr_dev_52.74.244.220;SSH;UNKNOWN;HARD;3;NRPE: Unable to read output
[1450873914] SERVICE ALERT: Nagios Server;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1450874751] SERVICE ALERT: credr_dev_52.74.244.220;SSH;OK;HARD;4;SSH OK - OpenSSH_6.6.1p1 Ubuntu-2ubuntu2 (protocol 2.0)
[1450875150] SERVICE ALERT: Nagios Server;Total Processes;OK;HARD;4;PROCS OK: 80 processes with STATE = RSZDT