[1455647531] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=50.26% system=1.03%, iowait=0.51%, idle=47.69%, nice=0.00%, steal=0.51%
[1455647591] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=50.25% system=0.00%, iowait=2.54%, idle=46.70%, nice=0.00%, steal=0.51%
[1455647651] SERVICE ALERT: credr_db_prod_middle_52.76.193.246;CPU Usage;OK;SOFT;3;CPU STATISTICS OK : user=0.51% system=1.01%, iowait=26.77%, idle=71.21%, nice=0.00%, steal=0.51%
[1455667209] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;1;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 135
[1455667269] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;SOFT;2;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 135
[1455667330] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;WARNING;HARD;3;POSTGRES_LOCKS WARNING: DB "apps_data" total locks: 127
[1455667869] SERVICE ALERT: credr_DW_52.76.17.124;PG Locks;OK;HARD;3;POSTGRES_LOCKS OK: DB "apps_data" total=28
[1455670901] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.77% system=3.98%, iowait=0.00%, idle=49.75%, nice=0.00%, steal=0.50%
[1455670961] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.50% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1455671021] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=44.72% system=5.03%, iowait=0.00%, idle=50.25%, nice=0.00%, steal=0.00%
[1455671381] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455674441] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=42.00% system=5.50%, iowait=2.00%, idle=50.00%, nice=0.00%, steal=0.50%
[1455674501] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=51.00% system=5.00%, iowait=0.50%, idle=43.00%, nice=0.00%, steal=0.50%
[1455674561] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=45.00% system=5.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455675101] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=28.50% system=0.50%, iowait=0.00%, idle=71.00%, nice=0.00%, steal=0.00%
[1455678159] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;CRITICAL;SOFT;1;CRITICAL -  Error: Redis connection to localhost:6379 failed - connect ECONNREFUSED
[1455678161] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=45.45% system=6.57%, iowait=0.00%, idle=47.98%, nice=0.00%, steal=0.00%
[1455678219] SERVICE ALERT: credr_stage_52.76.16.57;Redis Service;OK;SOFT;2;OK - uptime_in_days: 0
[1455678221] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;SOFT;2;CPU STATISTICS WARNING : user=44.50% system=6.00%, iowait=0.00%, idle=49.50%, nice=0.00%, steal=0.00%
[1455678281] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=46.00% system=4.00%, iowait=0.00%, idle=50.00%, nice=0.00%, steal=0.00%
[1455678821] SERVICE ALERT: credr_DW_52.76.17.124;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=30.65% system=2.01%, iowait=0.00%, idle=67.34%, nice=0.00%, steal=0.00%
[1455696761] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=83.50% system=2.00%, iowait=0.00%, idle=14.50%, nice=0.00%, steal=0.00%
[1455696821] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=2.51% system=0.50%, iowait=0.00%, idle=96.98%, nice=0.00%, steal=0.00%
[1455703661] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=62.30% system=1.57%, iowait=0.00%, idle=36.13%, nice=0.00%, steal=0.00%
[1455703721] SERVICE ALERT: credr_db_prod_master_52.76.120.118;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.50% system=1.00%, iowait=0.00%, idle=97.00%, nice=0.00%, steal=0.50%
[1455706181] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=78.50% system=20.00%, iowait=0.00%, idle=1.50%, nice=0.00%, steal=0.00%
[1455706241] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=76.88% system=21.61%, iowait=0.00%, idle=1.51%, nice=0.00%, steal=0.00%
[1455706301] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=79.60% system=18.91%, iowait=0.00%, idle=1.00%, nice=0.00%, steal=0.50%
[1455706481] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.50% system=0.00%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455712421] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=72.14% system=22.89%, iowait=0.00%, idle=4.98%, nice=0.00%, steal=0.00%
[1455712481] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=78.00% system=21.00%, iowait=0.00%, idle=1.00%, nice=0.00%, steal=0.00%
[1455712541] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=77.89% system=21.11%, iowait=0.00%, idle=1.01%, nice=0.00%, steal=0.00%
[1455712991] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=43.84% system=3.94%, iowait=1.97%, idle=49.75%, nice=0.00%, steal=0.49%
[1455713051] SERVICE ALERT: credr_prodsolr_52.74.114.132;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00%, iowait=0.00%, idle=100.00%, nice=0.00%, steal=0.00%
[1455714161] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;HARD;3;CPU STATISTICS WARNING : user=43.50% system=2.50%, iowait=0.00%, idle=53.50%, nice=0.00%, steal=0.50%
[1455714341] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;HARD;3;CPU STATISTICS OK : user=0.00% system=0.50%, iowait=0.00%, idle=99.50%, nice=0.00%, steal=0.00%
[1455731981] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;WARNING;SOFT;1;CPU STATISTICS WARNING : user=64.80% system=1.53%, iowait=0.00%, idle=33.16%, nice=0.00%, steal=0.51%
[1455732041] SERVICE ALERT: credr_stage_52.76.16.57;CPU Usage;OK;SOFT;2;CPU STATISTICS OK : user=1.01% system=0.00%, iowait=0.00%, idle=98.99%, nice=0.00%, steal=0.00%